The deubiquitinase Ubp3/Usp10 constrains glucose-mediated mitochondrial repression via phosphate budgeting

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

The study by Vengayil et al. presented a role for Ubp3 for mediating inorganic phosphate (Pi) compartmentalization in cytosol and mitochondria, which regulates metabolic flux between cytosolic glycolysis and mitochondrial processes. Although the exact function of increased Pi in mitochondria is not investigated, findings have valuable implications for understanding the metabolic interplay between glycolysis and respiration under glucose-rich conditions. They showed that UBP3 KO cells regulated decreased glycolytic flux by reducing the key Pi-dependent-glycolytic enzyme abundances, consequently increasing Pi compartmentalization to mitochondria. Increased mitochondria Pi increases oxygen consumption and mitochondrial membrane potential, indicative of increased oxidative phosphorylation. In conclusion, the authors reported that the Pi utilization by cytosolic glycolytic enzymes is a key process for mitochondrial repression under glucose conditions.

Comments on revised version:

This reviewer appreciates the author's responses addressing some of the concerns.

(1) However, the concern of reproducibility and experimental methods applied to the study is still valid, particularly considering that many conclusions were drawn from western blot analysis. The authors used separate gel loading controls for western blot analysis, which is not a valid method. Considering loading and other errors/discrepancies during the transfer phase of the assay, the direct control should be analyzing the membrane after transfer or using an internal control antibody on the same membrane. None of the western blots are indicated with marker sizes, and it isn't very clear how many repeats there are and whether those repeats are biological or technical repeats.

We thank the reviewer for raising this concern. This point requires detailed clarification regarding two key points: the first one regarding the use of Coomassie stained gels over internal ‘housekeeping gene’ antibodies, and the second one regarding the challenges in performing controls for western blots In case of high abundance proteins such as glycolytic enzymes.

(1) In our western blots, we have used Coomassie stained gel as a loading control for all our western blots. This is performed by cutting one half of the gel and using it for transfer followed by blotting and using the other half for Coomassie staining. I.e. This is not two separate gels that are loaded, but the same gel. Practically, this is no different from cutting a membrane to blot with different antibodies. This method is of course valid method for normalizing western blot data, and is used by multiple studies, for the reasons mentioned below. The historical use of a ‘house-keeping’ gene as a loading control for western blotting assumes that the protein levels of these does not change under different conditions. However, this approach has multiple, severe limitations (since a ‘housekeeping gene’ is entirely contextual, and indeed), and therefore it is correct to use total protein as a loading control. This is indeed recommended for use by multiple studies (Collins et al., 2015). Coomassie staining for total protein is far more reliable than using house-keeping genes as a loading control in western blots (Welinder and Ekblad, 2011). A notable example would be GAPDH itself, which is widely used as a loading control in many studies. As is clear from our data in this manuscript, GAPDH levels itself decrease in ubp3Δ cells. Had we used GAPDH as a loading control, we wouldn’t have identified the decrease in glycolytic enzymes in ubp3Δ cells, and this story would have met with a tragic fate very early on in its inception. We have in fact be very careful with these quantitations, and even before loading samples on gels, they are first normalized using a standard protein estimation assay (Bradford), followed by normalized loading, followed by cutting the gel into two parts - one for coomassie staining and protein normalization, and the other for the western blot for the respective proteins. However, in point (2) below, we clarify on why sometimes we have to load a separate gel with normalized protein, which should resolve this point.

(2) Glycolytic enzymes are highly abundant proteins and to achieve a signal in the linear range of western blot, the protein extracts have to be diluted (up to 25 or 50 times). As discussed under point 1, an internal control ‘housekeeping gene’ antibody is not a reliable method to use as loading control. Even if we want to use an antibody for an internal protein as a control, there are not many proteins that are as abundant as metabolic enzymes and because of this simple reason, the sample dilution results in these proteins not getting detected in the western blot since the signal will be below the limit of detection. This leaves using a separate gel loading control as the only easy to perform, reliable option.

We would like to further highlight the fact that the changes in metabolic enzymes and ETC proteins that we observe in the ubp3 mutant by western blot, were also independently observed by large scale untargeted quantitative proteomics study by (Isasa et al., 2015), which we cite extensively in this manuscript. Since an entirelyindependent study, using a completely different (untargeted) method has also shown very similar changes in proteins that we observe (mitochondrial, and glycolytic enzymes), there should be no room for doubt regarding the altered glycolytic enzyme and ETC protein levels that we discover in this study.

None of the western blots are indicated with marker sizes

We have clearly indicated the marker sizes in all our western blots. Separately, raw images of the blots and Coomassie stained gels have been provided with the manuscript raw data, and is therefore easily available for any interested reader.

It isn't very clear how many repeats there are and whether those repeats are biological or technical repeats.

We have already clearly indicated the details of each blot in the figure legends. For example “A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD.” We kindly request the reviewer to thoroughly go through the figure legends for details regarding the western blots, or any other data. We hope this addresses all the reviewer concerns regarding the credibility of our western blot results and the method of using Coomassie stained gels as loading controls in this study.

(2) Concern regarding citing the Ouyang et al. paper is still valid. This paper is an essential implication in phosphate metabolism and is directly related to some of the findings associated with mitochondrial function, along with conflicting results, which should be discussed in the discussion section. As a reviewer, I do not request citing any paper from the authors in general; however, considering some of the conflicting results here, citing and discussing paper from Ouyang et al. will improve the interoperation/value of their findings.

As mentioned in detail in our previous response letter, we do not believe that the study from Ouyang et al., present ‘conflicting results’ of any kind. Nevertheless, in response to the reviewer's suggestion, we have revised the discussion section of our manuscript and added a few points that incorporate the insights from Ouyang et al. These are in the discussion section (“It is important to highlight that our experiments, whether involving Pi supplementation or Pi limitations, maintain the cellular Pi concentration within the millimolar range and are conducted within a short timeframe (~ 1 hour). This differs significantly from Pi starvation studies, where cells are subjected to prolonged and complete Pi deprivation, triggering extensive metabolic adjustments to sustain available Pi pools, such as an increase in mitochondrial membrane potential, independent of respiration”). We trust that this modification will enhance the interested readers' understanding of our study's overarching conclusions.

Reviewer #2 (Public Review):

Summary:

Cells cultured in high glucose tend to repress mitochondrial biogenesis and activity, a prevailing phenotype type called Crabree effect that observed in different cell types and cancer. Many signaling pathways have been put forward to explain this effect. Vengayil et al proposed a new mechanism involved in Ubp3/Ubp10 and phosphate that controls the glucose repression of mitochondria. The central hypothesis is that ∆ubp3 shift the glycolysis to trehalose synthesis, therefore lead to the increase of Pi availability in the cytosol, then mitochondrial received more Pi and therefore the glucose repression is reduced.

Strengths:

The strength is that the authors used an array of different assays to test their hypothesis. Most assays were well designed and controlled.

Weaknesses:

I think the main conclusions are not strongly supported by the current dataset. Here are my comments on authors' response and model.

(1) The authors addressed some of my concerns related to ∆ubp3. But based on the results they observed and discussed, the ∆ubp3 redirect some glycolytic flux to gluconeogenesis while the 0.1% glucose in WT does not. Similarly, the shift of glycolysis to trehalose synthesis is also not relevant to the WT cells cultured in low glucose situation. This should be discussed in the manuscript to make sure readers are not misled to think ∆ubp3 mimic low glucose. It is likely that ∆ubp3 induce proteostasis stress, which is known to activate respiration and trehalose synthesis.

But based on the results they observed and discussed, the ∆ubp3 redirect some glycolytic flux to gluconeogenesis while the 0.1% glucose in WT does not. Similarly, the shift of glycolysis to trehalose synthesis is also not relevant to the WT cells cultured in low glucose situation.

We would like to clarify that we do not observe a redirection of glycolytic flux to gluconeogenesis in ubp3 mutant. What we observe is a rewiring of glycolytic flux into increased trehalose synthesis and PPP, and decreased glycolysis. Also, the shift of glycolysis to trehalose synthesis is relevant to WT cells cultured in low glucose. It is a well-known fact that the trehalose synthesis increases with decrease in media glucose. In case of 0.1% glucose, this increase in trehalose is not due to an increase in gluconeogenesis (since the pathways utilizing alternate carbon sources still remain repressed in 0.1% glucose (Yin et al., 2003)), but by the increase in glycolytic flux towards trehalose. This is also supported by increase in Tps2 protein levels upon decreasing glucose concentration (Shen et al., 2023). We will also note that there are very few studies that actually estimate gluconeogenic flux in cess (and they only rely on steady state measurements). Estimating gluconeogenic flux appropriately is challenging in itself (eg. see Niphadkar et al 2024).

In case of glucose concentrations lower than 0.1%, the shift to trehalose synthesis might not be as relevant. We observe that the glycolysis defective mutant tdh2tdh3 cells does not show an increase in trehalose synthesis (Figure 3-figure supplement 1E). However, in this context, the decrease in the rate of GAPDH catalyzed reaction alone appears to be sufficient to increase the Pi levels (Figure 3F) even without an increase in trehalose. Therefore, there might be differences in the relative contributions of these two arms towards Pi balance, based on whether it is low glucose in the environment, or a mutant such as ubp3Δ that modulates glycolytic flux. In ubp3Δ cells, the combination of low rate of GAPDH catalyzed reaction and high trehalose will happen (based on how glycolytic flux is modulated), vs only the low rate of the GAPDH catalyzed reaction in tdh2tdh3 cells. As an end point the increase in Pi happens in both cases, but this happens via slightly differing outcomes. Also note: in terms of free Pi sources a low-glucose condition (with low glycolytic rate) is very different from a no-glucose, respiratory condition (where cells perform very high gluconeogenesis, at a rate that is an order of magnitude higher than in low glucose). In respiration-reliant conditions such as in ethanol, cells switch to high gluconeogenesis, where there is a large increase in trehalose synthesis as a default (eg see Varahan et al 2019). In this condition, trehalose synthesis could become a major source for Pi (eg see Gupta 2021). This could also support the increased mitochondrial respiration. In an ethanol-only medium, the directionality of the GAPDH reaction is itself reversed (i.e. G-1,3-BP → G-3-P). Therefore, this reaction now becomes an added source of Pi, instead of a net consumer of Pi (see illustration in Figure 3G). Therefore, a very reasonable inference is that a combination of increased trehalose and increased 1,3 BPG to G3P conversion can become a Pi source, supporting increased mitochondrial respiration in a non-glucose, respiratory medium.

We have now clarified these points in the discussion section in the updated version of our manuscript. Lines xxx. We hope that this updated discussion section satisfies the reviewer’s concern regarding how relevant the increase in trehalose synthesis is for altered Pi balance and increased mitochondrial respiration in WT cells.

It is likely that ∆ubp3 induce proteostasis stress, which is known to activate respiration and trehalose synthesis.

Apart from some general changes in metabolism, there are no reports whatsoever that suggest that general proteostasis stress can results in an extensive, precise metabolic rewiring - where there is an increased in respiration, mitochondrial de-repression, precise decrease in two limiting glycolytic enzyme levels, and a precise reduction in glycolytic flux, as observed in the ubp3 mutant. If this was the case, deletion of any deubiquitinase should result in an increase in trehalose and respiration which clearly does not happen (as is already clear from the large screen shown in Figure 1)

However, in response to this query, we performed experiments to assess the extent of proteostasis stress in ubp3 mutants. For this, we have now estimated the changes in global ubiquitination in WT vs ubp3 mutant, and compared this with conditions of moderate proteostasis stress (mild heat shock at 42C/~1hr). These data are now included in the revised manuscript as Figure 1- figure supplement 1J. Notably, our analysis reveals only very minor alteration in global ubiquitination levels in ubp3 mutants compared to WT cells. This is in very stark contrast to limited heat stress, where a clear increase in global ubiquitination can be easily observed. Given these data, we can conclude that there is no significant general proteostatic stress in ubp3 mutants, that could induce substantial metabolic rewiring of such precise nature.

(2) Pi flux: it is known that vacuole can compensate the reduction of Pi in the cytosol. The paper they cited in the response, especially the Van Heerden et al., 2014 showed that the pulse addition of glucose caused transient Pi reduction and then it came back to normal level after 10min or so. If the authors mean the transient change of glycolysis and respiration, they should point that out clearly in the abstract and introduction. If the authors are trying to put out a general model, then the model must be reconsidered.

In Van Heerden et al., the pulse addition of glucose causes transient Pi reduction due to rapid Pi consumption in glycolysis. The phosphate levels came back to normal level because of the glucose flux into trehalose synthesis releasing free Pi. This is the entire crux of the study and this is the reason why tps2 mutants which cannot synthesize trehalose exhibit a growth defect and have decreased Pi levels. As explained in detail in our early response, the cellular Pi levels are maintained by a relative balance of reactions that consume and release Pi and therefore a change in this balance can change Pi as well. Indeed, if this were not the case, the tps2 mutants would simply maintain the Pi levels similar to WT cells by increasing Pi transport from the medium, which is clearly not the case (eg see Gupta 2021).

The cytosol has ~50mM Pi (van Eunen et al., 2010 FEBSJ), while only 1-2mM of glycolysis metabolites, not sure why partial reduction of several glycolysis enzymes will cause significant changes in cytosolic Pi level and make Pi the limiting factor for mitochondrial respiration. In response to this comment, the authors explained the metabolic flux that the rapid, continuous glycolysis will drain the Pi pool even each glycolytic metabolite is only 1-2mM. However, the metabolic flux both consume and release Pi, that's why there is such measurement of overall free Pi concentration amid the active metabolism. One possibility is that the observed cytosolic Pi level changes was caused by the measurement fluctuation.

The measurement fluctuations that we mentioned in our previous response letter was in case of cells grown in high and low glucose, where there are multiple factors such as mitochondrial amount which complicates the Pi measurements. In case of ubp3 mutants which have a similar amount of total mitochondria as that of WT cells, there is minimal fluctuation for Pi measurement. We have done extensive standardization of mitochondrial isolation and Pi measurement in the isolated mitochondria (as explained in detail in the manuscript) to minimize any such fluctuations.

However, the metabolic flux both consume and release Pi, that's why there is such measurement of overall free Pi concentration amid the active metabolism

The reviewer is correct in pointing out that metabolic flux consume and release Pi. However, in glucose grown yeast cells, the rate of glycolysis which is a Pi consuming reaction is higher than any other metabolic pathway. In fact, the glycolytic rate in glucose-grown S. cerevisiae is one of the highest ever observed in any living system. A decrease in glycolysis and an increase in trehalose therefore shifts the balance in Pi utilization and results in increased free Pi in ubp3 cells. For a more detailed theoretical reasoning on the consumption and production of Pi, see Gupta 2021.

Importantly, the authors measured Pi inside mito for ethanol and glucose, but not the cytosolic Pi, which is the key hypothesis in their model. The model here is that the glycolysis competes with mito for free cytosolic Pi, so it needs to inhibit glycolysis to free up cytosolic Pi for mitochondrial import to increase respiration. I don't see measurement of cytosolic Pi upon different conditions, only the total Pi or mito Pi. The fact is that in Fig.3C they saw WT+Pi in the medium increase total free Pi more than the ∆ubc3, while WT decrease mito Pi compared to WT control and ∆ubc3 and therefore decrease basal OCR upon Pi supplement. A simple math of Pitotal = Pi cyto + Pi mito tells us that if WT has more Pitotal (Fig.3C) but less Pi mito (fig.5 supp 1C), then it has higher Pi cyto. This is contradictory to what the authors tried to rationalize. Furthermore, as I pointed out previously, the isolated mitochondria can import more Pi when supplemented, so if there is indeed higher Picyto, then the mito in WT should import more Pi. So, to address these contradictory points, the authors must measure Pi in the cytosol, which is a critical experiment not done for their model. For example, they hypothesized that adding 2-DG, or ∆ubp3, suppress glycolysis and thus increase the supply of cytosolic Pi for mito to import, but no cytosolic Pi was measured (need absolute value, not the relative fold changes). It is also important to specific how the experiments are done, was the measurement done shortly after adding 2-DG. Given that the cells response to glucose changes/pulses differently in transient vs stable state, the authors are encouraged to specify that.

(1) Importantly, the authors measured Pi inside mito for ethanol and glucose, but not the cytosolic Pi, which is the key hypothesis in their model. The model here is that the glycolysis competes with mito for free cytosolic Pi, so it needs to inhibit glycolysis to free up cytosolic Pi for mitochondrial import to increase respiration. I don't see measurement of cytosolic Pi upon different conditions, only the total Pi or mito Pi.

As clearly described in the manuscript, the key hypothesis that emerges is the role of the availability/accessibility of Pi for the mitochondria, in the context of activity. As discussed in detail in the discussion section, this can come from a combination of available Pi pools in the cytosol and increased transport of this Pi to the mitochondria. While it is true that the decreased glycolysis in ubp3 mutants frees up available Pi pools in the cytosol, measurement of cytosolic Pi in these mutants growing in log phase might not necessarily show an increased cytosolic Pi, if the Pi is being actively transported the the mitochondria at a rate higher that the WT, as indicated by the ~6 fold increase in mitochondrial Pi in ubp3 cells. This would require tools such as intracellular fluorescence based-Pi sensors that could accurately capture temporal changes in cytosolic and mitochondrial Pi following glycolytic inhibition. However, these tools are not available till date for use in yeast and measuring cytosolic Pi following glycolytic inhibition over time using colorimetric Pi assays are extremely difficult.

However, the reviewer does correctly state that we had not included measurement of cytosolic Pi. Since the mitochondrial Pi estimate was itself a very challenging (and critical) experiment we had originally thought that data was sufficient. We have therefore now performed a series of new experiments, where we first enrich the cytosolic fraction (without mitochondrial contamination), and estimated cytosolic Pi amounts in WT and ubp3 cells. Our Pi measurements indicate a cytosolic Pi concentration in the range of ~35 mM, which is similar to the earlier reported values in yeast. We further observe that the cytosolic Pi is about ~25% lower in ubp3 mutants (~25-27 mM) compared to WT cells (Figure 4B). As mentioned earlier, this would be consistent with higher transport of Pi from the cytosol to the mitochondria in these cells. Effectively, ubp3 cells have a total increase in cellular Pi, and with a Pi pool distribution such that there is increased Pi availability in mitochondria (Figure 4B). This further substantiates this hypothesis of an increased Pi allocation to mitochondria in ubp3 mutants. The reason for increased rate of Pi transport to mitochondria is not immediately clear, but could also come from changes in cytosolic pH - a possibility that we suggest in our discussion, and is discussed in a later section of this response letter as well.

(2) The fact is that in Fig.3C they saw WT+Pi in the medium increase total free Pi more than the ∆ubc3, while WT decrease mito Pi compared to WT control and ∆ubc3 and therefore decrease basal OCR upon Pi supplement. A simple math of Pitotal = Pi cyto + Pi mito tells us that if WT has more Pitotal (Fig.3C) but less Pi mito (fig.5 supp 1C), then it has higher Pi cyto. This is contradictory to what the authors tried to rationalize. Furthermore, as I pointed out previously, the isolated mitochondria can import more Pi when supplemented, so if there is indeed higher Picyto, then the mito in WT should import more Pi.

a) “The fact is that in Fig.3C they saw WT+Pi in the medium increase total free Pi more than the ∆ubc3, while WT decrease mito Pi compared to WT control and ∆ubc3 and therefore decrease basal OCR upon Pi supplement. A simple math of Pitotal = Pi cyto + Pi mito tells us that if WT has more Pitotal (Fig.3C) but less Pi mito (fig.5 supp 1C), then it has higher Pi cyto.”

In WT cells supplemented with external Pi (WT+Pi), there is an increased total Pi, but a decreased mitochondrial Pi. As discussed in the discussion section in the manuscript, this could be due to the supplemented Pi not being transported to mitochondria. The reviewer is correct in pointing out that as per simple math this should mean that the cytosolic Pi in WT+Pi should be high. We have now assessed cytosolic Pi upon external Pi supplementation, and this is exactly what we observe in our cytosolic Pi measurements now included in the revised manuscript (Figure 5-figure supplement 5C). There is a higher cytosolic Pi in WT+Pi (~52 mM) compared to WT cells (~35 mM) and ubp3 cells (~27 mM). We have now pointed this out in the discussion section in the revised manuscript “Notably, this increased respiration does not happen upon direct Pi supplementation to highly glycolytic WT cells, where the Pi accumulates in cytosol, without increasing mitochondrial Pi (Figure 5-figure supplement 1C).” We hope that these new data completely addresses the reviewer’s concern regarding the Pi allocations in case of WT+Pi cells.

b) This is contradictory to what the authors tried to rationalize. Furthermore, as I pointed out previously, the isolated mitochondria can import more Pi when supplemented, so if there is indeed higher Picyto, then the mito in WT should import more Pi.

We would like to clarify that the Pi measurements in WT+Pi absolutely do not contradict our hypothesis. Furthermore, nowhere do we claim that an increase in cytosolic Pi will increase mitochondrial Pi!! On the contrary, we explain in detail that supplementing Pi to WT cells (which increases cytosolic Pi) will not increase respiration if the increased Pi is not being transported to mitochondria. This is exactly what happens in WT+Pi, where Pi accumulates in the cytosol but does not result in increased mitochondrial Pi. The reviewer argues that if there is higher cyto Pi, mitochondria should import more Pi. This is true in case of transport via diffusion where the external concentration dictates the direction of metabolite transport, but is fundamentally wrong in case of transport of metabolites where active transporters and additional regulators are involved. This is the entire basis of the idea of metabolic compartmentalisation where cells maintain pools of metabolites in different organelles which regulate the cellular metabolic state. A well-studied example is pyruvate, whose cytosolic concentration is high in glycolytic cells, but it's transport to mitochondria is reduced in glycolysis to maintain cytosolic fermentation. As discussed in the manuscript, a logical explanation for Pi supplementation not increasing respiration and mitochondria Pi is that there might be mechanisms in highly glycolytic cells that restrict the transport of Pi to mitochondria, thereby compartmentalizing Pi in the cytosol. One such possible mechanism is pH (discussed in a later section) and it is possible that there are other mechanisms involved.

In case of isolated mitochondria, Pi supplementation results in an increased respiration simply because it is an in vitro set up where we supplement metabolites such as pyruvate, malate and ADP along with phosphate to ensure that mitochondria is actively respiring and in this case Pi will be consumed since it is being used for ATP synthesis. This is entirely different from an in vivo scenario where cells are glycolytic, and mechanisms to prevent mitochondrial transport of metabolites such as pyruvate and phosphate are active.

c) It is also important to specific how the experiments are done, was the measurement done shortly after adding 2-DG?

Cells were treated with 2-DG for one hour and respiration was measured. We have mentioned these details clearly in the figure legends and methods.

d) The most likely model to me is that, which is also the consensus in the field, is that no matter 2-DG or ∆ubp3, the cells re-wiring metabolism in both cytosol and mitochondria, and it is the total network shift that cause the mitochondrial respiration increase, which requires the increase of mito import of Pi, ADP, O2, and substrates, but not caused/controlled by the Pi that singled out by the authors in their model.

The aim of our study is only to highlight the importance of mitochondrial Pi availability as a critical factor in controlling mitochondrial respiration. Of course this would require sufficient other factors such as ADP, substrates and oxygen. It cannot be otherwise. However, as we point out in the discussion, a major limiting factor might be Pi availability. While the altered glycolysis in ubp3 mutants might control availability of other factors such as pyruvate and ADP, this is not the focus of our study. We would also like to point out that prior studies show that even though cytosolic ADP decreases in the presence of glucose, this does not limit mitochondrial ADP uptake, or decrease respiration, due to the very high affinity of the mitochondrial ADP transporter. This is discussed in our discussion section as well. Further we show that the levels of ETC proteins can be altered by changing Pi levels, which places Pi as a major regulator of respiration. We would like to point out once again that studies in other systems have also highlighted a major role of mitochondrial Pi availability in controlling respiration. These references are included in our manuscript (Scheibye-Knudsen et al., 2009, Seifer et al., 2015). This includes a recent study in T cells that clearly shows increased mitochondrial respiration upon overexpressing mitochondrial Pi transporter SLC25A3 alone (Wu et al., 2023). Our manuscript now in fact provides a contextual explanation of these diverse observations from other cellular systems where mitochondrial Pi transport appears to regulate respiration.

(3) The explanation that cytosolic pH reduction upon glucose depletion/2DG is a mistake. There are a lot of data in the literature showing the opposite. If the authors do think this is true, then need to show the data. Again, it is important to distinguish transient vs stable state for pH changes.

We observe that directly supplementing Pi to WT cells growing in high glucose does not result in higher mitochondrial Pi or increased respiration. However, supplementing Pi to WT cells increases mitochondrial respiration in the presence of glycolytic inhibitor 2-DG. We therefore merely suggest that cytosolic pH could be an additional regulator of mitochondrial Pi transport, since this will be consistent with the differences in mitochondrial Pi transport in highly glycolytic cells, and cells with decreased glycolysis ( such as 2-DG addition and ubp3 mutant). This is because in mitochondria, Pi is co-transported along with protons. Therefore, changes in cytosolic pH (which changes the proton gradient) will control the mitochondrial Pi transport (Hamel et al., 2004). The glycolytic rate is itself a major factor that controls cytosolic pH. The cytosolic pH in highly glycolytic cells is maintained ~7, and decreasing glycolysis results in cytosolic acidification (Orij et al., 2011). Therefore, under conditions of decreased glycolysis (such as loss of Ubp3), cytosolic pH becomes acidic. Since mitochondrial Pi transport depends on the proton gradient, a low cytosolic pH would favour mitochondrial Pi transport. Therefore, under conditions of decreased glycolysis (2DG treatment, or loss of Ubp3), where cytosolic pH would be acidic, increasing cytosolic Pi might indirectly increase mitochondria Pi transport, thereby leading to increased respiration. But we certainly do leave alternate interpretations to the imagination of any reader, and are indeed open to them. These are all exciting future directions this study will enable a contextual interpretation of.

The explanation that cytosolic pH reduction upon glucose depletion/2DG is a mistake.

We have cited two independent studies which suggest that cytosolic pH decreases upon a decrease in glycolysis (Orij et al.,2011 ,Dechant et al., 2010). This control of cytosolic pH by the glycolytic rate has been extensively shown using glycolytic mutants, cells in low glucose and cells grown in the presence of glycolytic inhibitors. According to the reviewer, this is a mistake and

there are a lot of data in the literature showing the opposite.

In our literature review we did not come across any relevant studies that actually show the opposite. If the reviewer still thinks this is a mistake, the reviewer is welcome to include some of the relevant literature that clearly shows the opposite in the comments, with actual measurements of cytosolic pH. Additionally, the possible role of cytosolic pH in this context does not affect the conclusions of our study, and we only include this as a possibility in the discussion. Therefore, this is obviously well beyond the scope of experiments in our current study, and considering the extensive data from multiple studies that shows that cytosolic pH decreases under low glycolysis, there is no relevance to including experiments to address the same in this study. We leave this as a point for an interested reader to think about, and it certainly can nucleate new directions of future study.

Changes from the earlier version to the revised version

Abstract

Abstract

Many cells in high glucose repress mitochondrial respiration, as observed in the Crabtree and Warburg effects. Our understanding of biochemical constraints for mitochondrial activation is limited. Using a Saccharomyces cerevisiae screen, we identified the conserved deubiquitinase Ubp3 (Usp10), as necessary for mitochondrial repression. Ubp3 mutants have increased mitochondrial activity despite abundant glucose, along with decreased glycolytic enzymes, and a rewired glucose metabolic network with increased trehalose production. Utilizing Δubp3 cells, along with orthogonal approaches, we establish that the high glycolytic flux in glucose continuously consumes free Pi. This restricts mitochondrial access to inorganic phosphate (Pi), and prevents mitochondrial activation. Contrastingly, rewired glucose metabolism with enhanced trehalose production and reduced GAPDH (as in Δubp3 cells) restores Pi. This collectively results in increased mitochondrial Pi and derepression, while restricting mitochondrial Pi transport prevents activation. We therefore suggest that glycolytic-flux dependent intracellular Pi budgeting is a key constraint for mitochondrial repression.

Introduction

Earlier version: Rapidly proliferating cells have substantial metabolic and energy demands in order to increase biomass (1, 2). This includes a high ATP demand, obtained from cytosolic glycolysis or mitochondrial oxidative phosphorylation (OXPHOS), to fuel multiple reactions (3). Interestingly, many rapidly proliferating cells preferentially rely on ATP from glycolysis/fermentation over mitochondrial respiration even in oxygen-replete conditions, and is the well-known Warburg effect (4, 5). Many such cells repress mitochondrial processes in high glucose, termed glucose-mediated mitochondrial repression or the Crabtree effect (6, 7). This is observed in some tumors (5), neutrophils (8), activated macrophages (9), stem cells (10–12), and famously Saccharomyces cerevisiae (7). Numerous studies have identified signaling programs or regulators of glucose-mediated mitochondrial repression. However, biochemical programs and regulatory processes in biology evolve around key biochemical constraints (13). The biochemical constraints for mitochondrial repression remain unresolved (14, 15).

Revised version: Rapidly proliferating cells have substantial metabolic and energy demands in order to increase biomass (Cai and Tu, 2012; Zhu and Thompson, 2019). This includes a high ATP demand, obtained from cytosolic glycolysis or mitochondrial oxidative phosphorylation (OXPHOS), to fuel multiple reactions (Nelson et al., 2008). Interestingly, many rapidly proliferating cells preferentially rely on ATP from glycolysis/fermentation over mitochondrial respiration even in oxygen-replete conditions, and is the well-known Warburg effect (Vander Heiden et al., 2009; Warburg, 1925). Many such cells repress mitochondrial processes in high glucose, termed glucose-mediated mitochondrial repression or the Crabtree effect (Crabtree, 1929; De Deken, 1966). This is observed in some tumors (Vander Heiden et al., 2009), neutrophils (Xia et al., 2021), activated macrophages (Kornberg, 2020), stem cells (Abdel-Haleem et al., 2017; Pacini and Borziani, 2014; Tsogtbaatar et al., 2020), and famously Saccharomyces cerevisiae (De Deken, 1966). Numerous studies have identified signaling programs or regulators of glucose-mediated mitochondrial repression. However, biochemical programs and regulatory processes in biology evolve around key biochemical constraints (Cornish-Bowden, 2016). The biochemical constraints for mitochondrial repression remain unresolved (Diaz-Ruiz et al., 2011; Hammad et al., 2016).

Earlier version: There are two hypotheses on the biochemical principles driving mitochondrial repression. The first proposes direct roles for glycolytic intermediates in driving mitochondrial respiration, by inhibiting specific mitochondrial outputs (16, 17). The second hypothesizes that a competition between glycolytic and mitochondrial processes for mutually required metabolites/co-factors such as pyruvate, ADP or inorganic phosphate (Pi) could determine the extent of mitochondrial repression (14, 15, 18). These are not all mutually exclusive, and a combination of these factors might dictate mitochondrial repression in high glucose. However, any hierarchies of importance are unclear (19), and experimental data for the necessary constraints for mitochondrial repression remains incomplete.

Revised version: There are two hypotheses on the biochemical principles driving mitochondrial repression. The first proposes direct roles for glycolytic intermediates in driving mitochondrial respiration, by inhibiting specific mitochondrial outputs (Díaz-Ruiz et al., 2008; Lemus et al., 2018). The second hypothesizes that a competition between glycolytic and mitochondrial processes for mutually required metabolites/co-factors such as pyruvate, ADP or inorganic phosphate (Pi) could determine the extent of mitochondrial repression (Diaz-Ruiz et al., 2011; Hammad et al., 2016; Koobs, 1972). These are not all mutually exclusive, and a combination of these factors might dictate mitochondrial repression in high glucose. However, any hierarchies of importance are unclear (Rodríguez-Enríquez et al., 2001), and experimental data for the necessary constraints for mitochondrial repression remains incomplete.

Earlier version: One approach to resolve this question has been to identify regulators of metabolic state under high glucose. Post-translational modifications (PTMs) regulated by signaling systems can regulate mitochondrial repression (20–22). Ubiquitination is a PTM that regulates global proteostasis (23, 24), but the roles of ubiquitination-dependent processes in regulating mitochondrial repression are poorly explored. Ubiquitination itself is determined by the balance between ubiquitination, and deubiquitinase (DUB) dependent deubiquitination (25). Little is known about the roles of DUBs in regulating metabolic states, making the DUBs interesting candidate regulators of mitochondrial repression.

Revised version: One approach to resolve this question has been to identify regulators of metabolic state under high glucose. Post-translational modifications (PTMs) regulated by signaling systems can regulate mitochondrial repression (Broach, 2012; Hitosugi and Chen, 2014; Tripodi et al., 2015). Ubiquitination is a PTM that regulates global proteostasis (Hershko and Ciechanover, 1998; Komander and Rape, 2012), but the roles of ubiquitination-dependent processes in regulating mitochondrial repression are poorly explored. Ubiquitination itself is determined by the balance between ubiquitination, and deubiquitinase (DUB) dependent deubiquitination (Pickart and Eddins, 2004). Little is known about the roles of DUBs in regulating metabolic states, making the DUBs interesting candidate regulators of mitochondrial repression.

In this study, using an S. cerevisiae DUB knock-out library-based screen, we identified the evolutionarily conserved deubiquitinase Ubp3 (mammalian Usp10) as required for mitochondrial repression in high glucose. Loss of Ubp3 resulted in mitochondrial activation, along with a reduction in the glycolytic enzymes-phosphofructokinase 1 (Pfk1) and GAPDH (Tdh2 and Tdh3). This consequently reroutes glucose flux and increases trehalose biosynthesis. This metabolic rewiring increases Pi release from trehalose synthesis, and decreases Pi consumption in glycolysis, to cumulatively increase Pi pools. Using ubp3Δ cells along with independent analysis of wild-type cells, and isolated mitochondrial fractions, we establish that glycolytic-flux dependent Pi allocations to mitochondria determines mitochondrial activity. Through these data, we propose how intracellular Pi balance as controlled by glycolytic flux, is a biochemical constraint for mitochondrial repression.

Results

A deubiquitinase deletion screen identifies Ubp3 as a regulator of glucose-mediated mitochondrial repression

Earlier version: In cells such as S. cerevisiae, high glucose represses mitochondrial activity as well as OXPHOS-dependent ATP synthesis (7, 26) (illustrated in Figure 1A). Our initial objective was to identify proteostasis regulators of glucose-mediated mitochondrial repression. We generated and used a deubiquitinase (DUB) deletion strain library of S. cerevisiae (Figure 1-figure supplement 1A), to unbiasedly identify regulators of mitochondrial repression by measuring the fluorescence intensity of a potentiometric dye Mitotracker CMXRos (illustrated in Figure 1B). Using this screen, we identified DUB mutants with altered mitochondrial membrane potential (Figure 1C, Figure 1-figure supplement 1 A). Note: WT cells in a respiratory medium (2% ethanol) were used as a control to estimate maximum mitotracker fluorescence intensity (Figure 1-figure supplement 1B).

Revised version: In cells such as S. cerevisiae, high glucose represses mitochondrial activity as well as OXPHOS-dependent ATP synthesis (De Deken, 1966; Postma et al., 1989) (illustrated in Figure 1A). Our initial objective was to identify proteostasis regulators of glucose-mediated mitochondrial repression. We generated and used a deubiquitinase (DUB) deletion strain library of S. cerevisiae (Figure 1-figure supplement 1A), to unbiasedly identify regulators of mitochondrial repression by measuring the fluorescence intensity of a potentiometric dye Mitotracker CMXRos (illustrated in Figure 1B). Using this screen, we identified DUB mutants with altered mitochondrial membrane potential (Figure 1C, Figure 1-figure supplement 1 A). Note: WT cells in a respiratory medium (2% ethanol) were used as a control to estimate maximum mitotracker fluorescence intensity (Figure 1-figure supplement 1B).

A deubiquitinase deletion screen identifies Ubp3 as a regulator of glucose-mediated mitochondrial repression

A) Schematic depicting glucose-mediated mitochondrial repression (Crabtree effect).

B) Schematic describing the screen with a yeast DUB KO library to identify regulators of Crabtree effect.

C) Identifying DUB knockouts with altered mitochondrial potential. Heat map shows relative mitochondrial membrane potential of 19 deubiquitinase deletions in high glucose, from 2 biological replicates. Also see Figure 1-figure supplements 1A, B.

D) The deubiquitinase activity of Ubp3 and repression of mitochondrial membrane potential. WT, ubp3Δ, and Ubp3C469A were grown inhigh glucose and relative mitochondrial membrane potential was measured. Data represent mean ± SD from three biological replicates (n=3). Also see Figure 1-figure supplement 1 D.

E) Effect of loss of Ubp3 on ETC complex IV subunit Cox2. WT and ubp3Δ were grown in high glucose, and Cox2 was measured (western blot using an anti-Cox2 antibody). A representative blot (out of 3 biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD.

F) Basal oxygen consumption rate (OCR) in high glucose in ubp3Δ. WT and ubp3Δ were grown in high glucose, and OCR was measured. Basal OCR corresponding to ∼3 × 10^5 cells, from two independent experiments (n=2), normalized to the OD600 is shown. Bar graph representations are shown in the inset. Data represent mean ± SD.

G) Total ATP levels in ubp3Δ and WT. WT and ubp3Δ were grown in high glucose, and total ATP were measured. Data represent mean ± SD from three biological replicates (n=3).

H) Dependence of ubp3Δ on mitochondrial ATP. WT and ubp3Δ cells were grown in high glucose, and treated with 1 mM sodium azide for 45 minutes. Total ATP levels in sodium azide treated and untreated cells were measured. Data represent mean ± SD (n=3).

I) Requirement for mitochondrial respiration in high glucose in ubp3Δ. A growth curve of WT and ubp3Δ in high glucose in the presence of OXPHOS uncoupler FCCP (10 µM), and serial dilution growth assay in high glucose in the presence/absence of sodium azide (1 mM) are shown. Data represent mean ± SD (n=2). Also see Figure 1-figure supplement 1 H, I.

Data information: **p<0.01, ***p<0.001.

Earlier version: A prominent ‘hit’ was the evolutionarily conserved deubiquitinase Ubp3 (Figure 1C), (homologous to mammalian Usp10) (Figure 1-figure supplement 1D). Due to its high degree of conservation across eukaryotes (Figure 1-figure supplement 1D) as well as putative roles in metabolism or mitochondrial function (27–29), we focused our further attention on this DUB. Cells lacking Ubp3 showed an ∼1.5-fold increase in mitotracker fluorescence (Figure 1C, Figure 1-figure supplement 1A, C). Cells with catalyticaly inactive Ubp3 (Ubp3C469A), showed increased mitochondrial potential comparable to ubp3Δ (Figure 1D). This data confirmed that Ubp3 catalytic activity of is required to fully repress mitochondrial activity under high glucose. The catalytic site mutation did not affect steady-state Ubp3 levels (Figure 1-figure supplement 1E).

Revised version: A prominent ‘hit’ was the evolutionarily conserved deubiquitinase Ubp3 (Figure 1C), (homologous to mammalian Usp10) (Figure 1-figure supplement 1D). Due to its high degree of conservation across eukaryotes (Figure 1-figure supplement 1D) as well as putative roles in metabolism or mitochondrial function (Isasa et al., 2015; Nostramo et al., 2016; Ossareh-Nazari et al., 2010), we focused our further attention on this DUB. Cells lacking Ubp3 showed an ∼1.5-fold increase in mitotracker fluorescence (Figure 1C, Figure 1-figure supplement 1A, C). Cells with catalyticaly inactive Ubp3 (Ubp3C469A), showed increased mitochondrial potential comparable to ubp3Δ (Figure 1D). This data confirmed that Ubp3 catalytic activity of is required to fully repress mitochondrial activity under high glucose. The catalytic site mutation did not affect steady-state Ubp3 levels (Figure 1-figure supplement 1E).

Earlier version: Next, to assess the requirement of Ubp3 for mitochondrial function, we quantified the electron transport chain (ETC) complex IV subunit Cox2 (30)(Figure 1E). ubp3Δ had higher Cox2 than WT (Figure 1E). As a control, we estimated total mitochondrial content in WT and ubp3Δ cells, using either estimates of the structural protein Tom70, or measuring the fluorescence-intensity in strains engineered with mitochondrial targeted mNeonGreen (31). There was no increase in the total mitochondrial volume (estimated by measuring the intensity of mitochondria targeted mNeon green) (Figure 1-figure supplement 1F) or mitochondrial outer membrane protein Tom70 (Figure 1-figure supplement 1G). This suggests that the increased Cox2 is not merely because of higher total mitochondrial content. We next measured the basal oxygen consumption rate (OCR) of ubp3Δ, and basal OCR was higher in ubp3Δ, indicating higher respiration (Figure 1F).

Revised version: Next, to assess the requirement of Ubp3 for mitochondrial function, we quantified the electron transport chain (ETC) complex IV subunit Cox2 (Fontanesi et al., 2006)(Figure 1E). ubp3Δ had higher Cox2 than WT (Figure 1E). As a control, we estimated total mitochondrial content in WT and ubp3Δ cells, using either estimates of the structural protein Tom70, or measuring the fluorescence-intensity in strains engineered with mitochondrial targeted mNeonGreen (Dua et al., 2022). There was no increase in the total mitochondrial volume (estimated by measuring the intensity of mitochondria targeted mNeon green) (Figure 1-figure supplement 1F) or mitochondrial outer membrane protein Tom70 (Figure 1-figure supplement 1G). This suggests that the increased Cox2 is not merely because of higher total mitochondrial content. We next measured the basal oxygen consumption rate (OCR) of ubp3Δ, and basal OCR was higher in ubp3Δ, indicating higher respiration (Figure 1F).

Next, we asked if mitochondrial ATP synthesis was higher in ubp3Δ. The total ATP levels in WT and ubp3Δ were comparable (Figure 1G). However, upon treatment with the ETC complex IV inhibitor sodium azide, ATP levels in WT were higher than ubp3Δ, contributing to ∼60% of the total ATP (Figure 1H). In contrast, the ATP levels in ubp3Δ after sodium azide treatment were ∼ 40% of the total ATP (Figure 1H). These data suggest a higher contribution of mitochondrial ATP synthesis towards the total ATP pool in ubp3Δ.

Earlier version: We next asked if ubp3Δ required higher mitochondrial activity for growth, using a series of mitochondrial activity inhibitors and comparing relative growth. In high glucose, WT cells show minimal growth inhibition in the presence of sodium azide, indicating lower reliance on mitochondrial function (Figure 1I). Contrastingly, ubp3Δ or Ubp3C469A show a severe growth defect in the presence of the mitochondrial ETC complex inhbitors sodium azide and oligomycin, and the mitochondrial OXPHOS uncoupler FCCP (Figure 1I, Figure 1-figure supplement 1H). Additionally, the loss of Ubp3 in respiration defective cox2-62 cells (32), or Rho0 cells (which lacks mitochondrial DNA) resulted in a severe growth defect (Figure 1-figure supplement 1I). Deletion of ATP synthase subunits Atp1 and Atp10 also results in a severe growth defect in ubp3Δ compared to WT (Figure 1-figure supplement 1I). Together, these results indicate that the loss of Ubp3 makes cells dependent on mitochondrial ATP synthesis in high glucose.

Revised version: We next asked if ubp3Δ required higher mitochondrial activity for growth, using a series of mitochondrial activity inhibitors and comparing relative growth. In high glucose, WT cells show minimal growth inhibition in the presence of sodium azide, indicating lower reliance on mitochondrial function (Figure 1I). Contrastingly, ubp3Δ or Ubp3C469A show a severe growth defect in the presence of the mitochondrial ETC complex inhbitors sodium azide and oligomycin, and the mitochondrial OXPHOS uncoupler FCCP (Figure 1I, Figure 1-figure supplement 1H). Additionally, the loss of Ubp3 in respiration defective cox2-62 cells (Bonnefoy et al., 2001), or Rho0 cells (which lacks mitochondrial DNA) resulted in a severe growth defect (Figure 1-figure supplement 1I). Deletion of ATP synthase subunits Atp1 and Atp10 also results in a severe growth defect in ubp3Δ compared to WT (Figure 1-figure supplement 1I). Together, these results indicate that the loss of Ubp3 makes cells dependent on mitochondrial ATP synthesis in high glucose.

Added in the revised version: In order to address whether the deletion of Ubp3 might increase ubiquitinated proteins and consequent proteostatic stress, we analysed the global ubquitination state in ubp3Δ cells. WT and ubp3Δ cells grown under brief (1 hour) heat stress, which increases protein ubiquitination, was used as a control. However, we did not observe a significant increase in the global ubiqutinatin state in ubp3Δ cells (Figure 1-figure supplement 1J). This suggests that the altered mitochondrial metabolism in ubp3Δ cells is unlikely to be due to proteostatic stress.

Collectively these data show that in 2% glucose, ubp3Δ have high mitochondrial activity, respiration, and rely on this mitochondrial function for ATP production and growth. We therefore decided to use ubp3Δ cells to start delineating requirements for glucose-mediated mitochondrial repression.

Key glycolytic enzymes decrease and glucose flux is rerouted in ubp3Δ cells

Earlier version: Glucose-6 phosphate (G6P) is the central node in glucose metabolism where carbon-allocations are made towards distinct metabolic arms, primarily: glycolysis, the pentose phosphate pathway (PPP), and trehalose biosynthesis (Figure 2A). We first compared amounts of two key (‘rate-controling’) glycolytic enzymes - Phosphofructokinase 1 (Pfk1), GAPDH isozymes (Tdh2, Tdh3) (3), along with the enolase isozymes (Eno1, Eno2) in WT and ubp3Δ cells. Pfk1, Tdh2, and Tdh3 substantially decreased in ubp3Δ (but not Eno1 and Eno2) (Figure 2B, Figure 2-figure supplement 1A). Since deubiquitinases can control protein amounts by regulating proteasomal degradation, we asked if the decrease in Pfk1, Tdh2 and Tdh3 in ubp3Δ is due to proteasomal degradation. To test this, we measured the levels of these enzymes in ubp3Δ after treatement with proteasomal inhibitor MG132. We did not observe any rescue in the levels of these enzymes in MG132 treated samples, suggesting that the decreased levels of these enzymes were not due to increased proteasomal degradation (Figure 2-figure supplement 1B). To further understand if these enzyme transcripts are altered in ubp3Δ, we measured the mRNA levels of PFK1, TDH2 and TDH3 in WT, ubp3Δ and Ubp3C469A cells. The transcripts of all the three genes in ubp3Δ and Ubp3C469A cells decreased (Figure 2-figure supplement 1C), suggesting that Ubp3 regulates the transcripts of these glycolytic enzyme genes. The reduction in the Pfk and GAPDH enzyme amounts was intriguing, because the Pfk and GAPDH steps are critical in determining glycolytic flux (3, 33). A reduction in these enzymes could therefore decrease glycolytic flux, and would reroute glucose (G6P) allocations via mass action towards other branches of glucose metabolism, primarily the pentose phosphate pathway as well as trehalose biosynthesis (Figure 2A). To assess this, we first measured the steady-state levels of key glycolytic and pentose phosphate pathway (PPP) intermediates, and trehalose in WT or ubp3Δ using targeted LC-MS/MS (Figure 2C, Figure 2-figure supplement 1D). Glucose-6/fructose-6 phosphate (G6P/F6P) increased in ubp3Δ (Figure 2C). Concurrently, trehalose, and the PPP intermediates ribose 5-phosphate (R5P) and sedoheptulose 7-phosphate (S7P), increased in ubp3Δ (Figure 2C).

Revised version: Glucose-6 phosphate (G6P) is the central node in glucose metabolism where carbon-allocations are made towards distinct metabolic arms, primarily: glycolysis, the pentose phosphate pathway (PPP), and trehalose biosynthesis (Figure 2A). We first compared amounts of two key (‘rate-controling’) glycolytic enzymes - Phosphofructokinase 1 (Pfk1), GAPDH isozymes (Tdh2, Tdh3) (Nelson et al., 2008), along with the enolase isozymes (Eno1, Eno2) in WT and ubp3Δ cells. Pfk1, Tdh2, and Tdh3 substantially decreased in ubp3Δ (but not Eno1 and Eno2) (Figure 2B, Figure 2-figure supplement 1A). Since deubiquitinases can control protein amounts by regulating proteasomal degradation, we asked if the decrease in Pfk1, Tdh2 and Tdh3 in ubp3Δ is due to proteasomal degradation. To test this, we measured the levels of these enzymes in ubp3Δ after treatement with proteasomal inhibitor MG132. We did not observe any rescue in the levels of these enzymes in MG132 treated samples, suggesting that the decreased levels of these enzymes were not due to increased proteasomal degradation (Figure 2-figure supplement 1B). To further understand if these enzyme transcripts are altered in ubp3Δ, we measured the mRNA levels of PFK1, TDH2 and TDH3 in WT, ubp3Δ and Ubp3C469A cells. The transcripts of all the three genes in ubp3Δ and Ubp3C469A cells decreased (Figure 2-figure supplement 1C), suggesting that Ubp3 regulates the transcripts of these glycolytic enzyme genes. The reduction in the Pfk and GAPDH enzyme amounts was intriguing, because the Pfk and GAPDH steps are critical in determining glycolytic flux (Nelson et al., 2008; Shestov et al., 2014). A reduction in these enzymes could therefore decrease glycolytic flux, and would reroute glucose (G6P) allocations via mass action towards other branches of glucose metabolism, primarily the pentose phosphate pathway as well as trehalose biosynthesis (Figure 2A). To assess this, we first measured the steady-state levels of key glycolytic and pentose phosphate pathway (PPP) intermediates, and trehalose in WT or ubp3Δ using targeted LC-MS/MS (Figure 2C, Figure 2-figure supplement 1D). Glucose-6/fructose-6 phosphate (G6P/F6P) increased in ubp3Δ (Figure 2C). Concurrently, trehalose, and the PPP intermediates ribose 5-phosphate (R5P) and sedoheptulose 7-phosphate (S7P), increased in ubp3Δ (Figure 2C).

Key glycolytic enzymes decrease and glucose flux is rerouted in ubp3Δ cells

A) A schematic illustrating directions of glucose-6-phosphate (G6) flux in cells. Glucose is converted to G6P, a precursor for trehalose, the pentose phosphate pathway (PPP), and glycolysis.

B) Effect of loss of Ubp3 on key glycolytic enzymes. WT and ubp3Δ were grown in high glucose and the Pfk1, Tdh2, and Tdh3 levels were measured by western blot using an anti-FLAG antibody. A representative blot (out of three biological replicates, n=3) and their quantification are shown. Data represent mean ± SD. Also see Figure 2-figure supplement 1A.

C) Steady-state metabolite amounts in WT and ubp3Δ in high glucose. Relative steady-state levels of trehalose, major glycolytic, and PPP intermediates were estimated in WT and ubp3Δ. Data represent mean ± SD from three biological replicates (n=3). Also see Appendix Table S3.

D) Relative glycolytic and trehalose synthesis flux in WT and ubp3Δ. Relative 13C-label incorporation into trehalose and glycolytic intermediates, after a pulse of 1% 13C6 glucose is shown. Data represent mean ± SD from three biological replicates (n=3). Also see Appendix Table S3, Figure 2-figure supplement 1 D, E.

E) Ethanol production in ubp3Δ. WT and ubp3Δ were grown in high glucose and ethanol in the media was measured. Data represent mean ± SD from three biological replicates (n=3).

F) Relative rate of ethanol production in WT vs ubp3Δ. WT and ubp3Δ were grown in high glucose (to OD600 ∼ 0.6), equal numbers of cells were shifted to fresh medium (high glucose) and ethanol concentration in the medium was measured temporally. Data represent mean ± SD from three biological replicates (n=3)

Data information: *p<0.05, **p<0.01, ***p<0.001.

Since steady-state metabolite amounts cannot separate production from utilization, in order to unambiguously assess if glycolytic flux is reduced in ubp3Δ cells, we utilized a pulse labeling of 13C6 glucose, following which the label incorporation into glycolytic and other intermediates was measured. Note that because glycolytic flux is very high in yeast, this experiment would require rapid pulsing and extraction of metabolites in order to stay in a linear range and avoid label saturation. We therefore established a very short time point of label-addition, quenching and metabolite extraction post 13C glucose pulse. Since flux saturates/reaches steady-state in seconds, we first ensured that the label incorporation into individual metabolites after the 13C glucose pulse was in the linear range, and for early glycolytic intermediates this was seconds after glucose addition. This new methodology is extensively described in materials and methods, with required controls shown in Figure 2-figure supplement 1F. WT and ubp3Δ cells were grown in high glucose, pulsed with 13C6 glucose, and the relative 13C label incorporation into glycolytic intermediates and trehalose were measured, as shown in the schematic (Figure 2-figure supplement 1E). In ubp3Δ, 13C label incorporation into G6P/F6P as well as trehalose substantially increased (Figure 2D). Contrastingly, 13C label incorporation into glycolytic intermediates F1,6BP, G3P, 3PG and PEP decreased, indicating decreased glycolytic flux (Figure 2D). We next measured ethanol concentrations and production rates, as an additional output of relative glycolytic rates. We observed decreased steady-state ethanol levels, as well as ethanol production rates in ubp3Δ (Figure 2F).

Glycolysis-derived pyruvate is transported to mitochondria and fuels the TCA cycle. Therefore, we asked if the decreased glycolytic rate result in a decrease in the TCA cycle flux as well. To test this, we first measured the steady state levels of TCA cycle intermediates in WT or ubp3Δ using targeted LC-MS/MS. We did not observe any significant change in the levels of TCA cycle intermediates in ubp3Δ, except malate which showed a significant decrease in ubp3Δ (Figure 2-figure supplement 2A). Next, in order to assess if TCA cycle flux reduces in ubp3Δ cells, WT and ubp3Δ cells were grown in high glucose, pulsed with 13C6 glucose, and the relative 13C label incorporation into TCA cycle intermediates was measured, as shown in the schematic (Figure 2-figure supplement 2B). The kinetics of 13C label incorporation in TCA cycle intermediates are shown in Figure 2-figure supplement 2C. We did not observe any significant change in the relative 13C label incorporation in TCA cycle intermediates in ubp3Δ. Therefore, these data suggest that the decreased glycolytic flux in ubp3Δ does not result in a decrease in TCA cycle flux. The increased respiration and mitochondrial activity in ubp3Δ cells is therefore driven via other factors.

Collectively, these results reveal that that reduced Pfk1 and GAPDH (in ubp3Δ) decrease glucose flux via glycolysis, which results in rewired glucose flux towards trehalose biosynthesis and the PPP.

Rerouted glucose flux results in phosphate (Pi) accumulation

Earlier version: We therefore asked if the proteomic state, as observed in ubp3Δ cells, could provide clues to explain the coupling between mitochondrial derepression and rerouted glucose flux. A recent study by Isasa et al., had systematically quantified the changes in protein levels in ubp3Δ cells (28). We therefore reanalyzed this extensive dataset, looking for changes in proteins that would correlate with these metabolic processes. Notably, we observed increased levels of proteins of the mitochondrial ETC and respiration, and decreased amounts of glucose metabolizing enzymes in ubp3Δ (Figure 3A). Additionally, multiple proteins involved in regulating phosphate (Pi) homeostasis were decreased in ubp3Δ (Figure 3A) (28). We had recently uncovered a reciprocal coupling of Pi homeostasis with the different arms of glucose metabolism, particularly trehalose biosynthesis (34, 35), and therefore wondered if a glycolytic-flux dependent change in Pi homeostatsis had any role in mitochondrial respiration. Our hypothesis was refined based on the reasoning given below.

Revised version: We therefore asked if the proteomic state, as observed in ubp3Δ cells, could provide clues to explain the coupling between mitochondrial derepression and rerouted glucose flux. A recent study by Isasa et al., had systematically quantified the changes in protein levels in ubp3Δ cells (Isasa et al., 2015). We therefore reanalyzed this extensive dataset, looking for changes in proteins that would correlate with these metabolic processes. Notably, we observed increased levels of proteins of the mitochondrial ETC and respiration, and decreased amounts of glucose metabolizing enzymes in ubp3Δ (Figure 3A). Additionally, multiple proteins involved in regulating phosphate (Pi) homeostasis were decreased in ubp3Δ (Figure 3A) (Isasa et al., 2015). We had recently uncovered a reciprocal coupling of Pi homeostasis with the different arms of glucose metabolism, particularly trehalose biosynthesis (Gupta et al., 2019; Gupta and Laxman, 2021), and therefore wondered if a glycolytic-flux dependent change in Pi homeostatsis had any role in mitochondrial respiration. Our hypothesis was refined based on the reasoning given below.

Rerouted glucose flux results in phosphate (Pi) accumulation

A) Changes in protein levels in ubp3Δ (dataset from Isasa et al., 2015). ubp3Δ cells have an increase in proteins involved in mitochondrial respiration and decrease in proteins involved in glucose and phosphate metabolism.

B) Schematic showing maintenance of Pi balance during glycolysis. Trehalose synthesis from G6P releases Pi, and the conversion of G3P to 1,3BPG by GAPDH consumes Pi. In ubp3Δ, trehalose biosynthesis (which releases Pi) increases. ubp3Δ have decreased GAPDH, which will decrease Pi consumption. This increase in Pi release along with decreased Pi consumption could increase cytosolic Pi.

C) Intracellular Pi levels in WT and ubp3Δ. WT and ubp3Δ were grown in high glucose and the total free phosphate (Pi) levels were estimated. WT in high Pi (2% glucose, 10mM Pi) was a positive control. Data represent mean ± SD from three biological replicates (n=3). Also see Figure 3-figure supplement 1A.

D) Pho regulon responses in WT and ubp3Δ. Protein levels of Pho84-FLAG and Pho12-FLAG were compared between WT grown in high glucose and in high Pi, ubp3Δ in high glucose with or without a shift to a no-Pi medium for one hour, by western blot. A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD.

E) Contribution of trehalose synthesis as a Pi source. WT, tps2Δ, ubp3Δ, and ubp3Δtps2Δ were grown in high glucose and the total Pi levels were estimated. Data represent mean ± SD from three biological replicates (n=3). Also see Figure 3-figure supplement 1B.

F) Loss of GAPDH isozymes Tdh2 and Tdh3 and effect on Pi. WT, ubp3Δ, and tdh2Δtdh3Δ were grown in high glucose and total Pi was estimated. Data represent mean ± SD from three biological replicates (n=3).

G) Pi levels in ubp3Δ and tdh2Δtdh3Δ cells in ethanol medium. WT, ubp3Δ, and tdh2Δtdh3Δ cells were grown in ethanol medium and the total Pi levels were estimated. Data represent mean ± SD from three biological replicates (n=3).

Data information: *p<0.05, **p<0.01, ***p<0.001, ****p<0.0001.

Earlier version: One explanation for mitochondrial repression can be an internal competition for shared metabolites/co-factors between (cytosolic) glycolytic and mitochondrial processes, that mitochondria might not be sufficiently able to access (14, 18, 19). In this context, a plausible role for inorganic phosphate (Pi) in regulating mitochondrial repression can be hypothesized. Cytosolic glycolysis requires rapid, high consumption of net Pi (19, 36, 37), and this could possibly limit the Pi that is continuously available for mitochondrial use (18, 38), thereby repressing mitochondria. Contextually, the balance between reactions releasing vs consuming Pi could explain changes in global Pi levels (35). A continuous hub of Pi consumption is glycolysis. Here, GAPDH catalyzes G3P to 1,3BPG, converting ADP to ATP, while concurrently consuming a molecule of Pi (37, 39). This Pi that goes into ATP will subsequently be used for nucleotide biosynthesis, polyphosphate biosynthesis and protein phosphorylation (35, 40). Therefore, we can surmise that in high glycolytic flux, the production of ATP, nucleotides and polyphosphates is concurrent with Pi consumption (39, 41, 42). Could this reaction therefore limit cytosolic Pi for the mitochondria (as illustrated in Figure 3B)? Notably, ubp3Δ have reduced GAPDH levels and decreased glycolytic flux. Second, trehalose synthesis is a Pi releasing reaction, and a major source of free Pi that is critical for Pi homeostasis (34, 37). Flux through this reaction is also substantially higher in ubp3Δ. Therefore, these cells might have increased Pi release (via trehalose), coupled with decreased Pi consumption (via decreased GAPDH). We therefore asked if total Pi increases in ubp3Δ (Figure 3B)?

Revised version: One explanation for mitochondrial repression can be an internal competition for shared metabolites/co-factors between (cytosolic) glycolytic and mitochondrial processes, that mitochondria might not be sufficiently able to access (Diaz-Ruiz et al., 2011; Koobs, 1972; Rodríguez-Enríquez et al., 2001). In this context, a plausible role for inorganic phosphate (Pi) in regulating mitochondrial repression can be hypothesized. Cytosolic glycolysis requires rapid, high consumption of net Pi (Mason et al., 1981; Rodríguez-Enríquez et al., 2001; Van Heerden et al., 2014), and this could possibly limit the Pi that is continuously available for mitochondrial use (Brazy et al., 1982; Koobs, 1972), thereby repressing mitochondria. Contextually, the balance between reactions releasing vs consuming Pi could explain changes in global Pi levels (Gupta and Laxman, 2021). A continuous hub of Pi consumption is glycolysis. Here, GAPDH catalyzes G3P to 1,3BPG, converting ADP to ATP, while concurrently consuming a molecule of Pi (Hohmann et al., 1996; Van Heerden et al., 2014). This Pi that goes into ATP will subsequently be used for nucleotide biosynthesis, polyphosphate biosynthesis and protein phosphorylation (Gupta and Laxman, 2021; Hunter, 2012). Therefore, we can surmise that in high glycolytic flux, the production of ATP, nucleotides and polyphosphates is concurrent with Pi consumption (Austin and Mayer, 2020; Hohmann et al., 1996; Ljungdahl and Daignan-Fornier, 2012). Could this reaction therefore limit cytosolic Pi for the mitochondria (as illustrated in Figure 3B)? Notably, ubp3Δ have reduced GAPDH levels and decreased glycolytic flux. Second, trehalose synthesis is a Pi releasing reaction, and a major source of free Pi that is critical for Pi homeostasis (Gupta et al., 2019; Van Heerden et al., 2014). Flux through this reaction is also substantially higher in ubp3Δ. Therefore, these cells might have increased Pi release (via trehalose), coupled with decreased Pi consumption (via decreased GAPDH). We therefore asked if total Pi increases in ubp3Δ (Figure 3B)?

Earlier version: To test this, we directly assessed total Pi levels in ubp3Δ and WT. ubp3Δ cells had higher Pi than WT, and this was comparable to Pi in WT grown in excess Pi (Figure 3C). Similarly, Pi amounts also increased in Ubp3C469A (Figure 3-figure supplement 1A). Therefore, the loss of Ubp3 increases intracellular Pi levels. Next, we asked if ubp3Δ cells exhibit signatures of a ‘high Pi’ state. S. cerevisiae maintains internal Pi balance by controlling the expression of multiple genes collectively known as the Pho regulon (43). The Pho regulon is induced under Pi limitation, and repressed during Pi sufficiency (34, 43). We assessed two major Pho proteins (Pho84: a high-affinity membrane Pi transporter, and Pho12: an acid phosphatase) in WT and ubp3Δ in high glucose. ubp3Δ cells have lower amounts of Pho84 and Pho12 (Figure 3D). Further, upon shifting to low Pi for one hour, Pho84 and Pho12 increased in ubp3Δ. These data suggest that reduced Pho84 and Pho12 amounts in ubp3Δ are because of increased Pi, and not due to altered Pho regulon function itself (Figure 3D). These data clarify earlier observations from ubp3Δ which noted reduced Pho proteins (28). Therefore, ubp3Δ constitutively have higher Pi, and likely a consequent decrease in Pho proteins.

Revised version: To test this, we directly assessed total Pi levels in ubp3Δ and WT. ubp3Δ cells had higher Pi than WT, and this was comparable to Pi in WT grown in excess Pi (Figure 3C). Similarly, Pi amounts also increased in Ubp3C469A (Figure 3-figure supplement 1A). Therefore, the loss of Ubp3 increases intracellular Pi levels. Next, we asked if ubp3Δ cells exhibit signatures of a ‘high Pi’ state. S. cerevisiae maintains internal Pi balance by controlling the expression of multiple genes collectively known as the Pho regulon (Mouillon and Persson, 2006). The Pho regulon is induced under Pi limitation, and repressed during Pi sufficiency (Gupta et al., 2019; Mouillon and Persson, 2006). We assessed two major Pho proteins (Pho84: a high-affinity membrane Pi transporter, and Pho12: an acid phosphatase) in WT and ubp3Δ in high glucose. ubp3Δ cells have lower amounts of Pho84 and Pho12 (Figure 3D). Further, upon shifting to low Pi for one hour, Pho84 and Pho12 increased in ubp3Δ. These data suggest that reduced Pho84 and Pho12 amounts in ubp3Δ are because of increased Pi, and not due to altered Pho regulon function itself (Figure 3D). These data clarify earlier observations from ubp3Δ which noted reduced Pho proteins (Isasa et al., 2015). Therefore, ubp3Δ constitutively have higher Pi, and likely a consequent decrease in Pho proteins.

Earlier version: We next asked if the increased Pi in ubp3Δ is because of altered G6P allocations towards different end-points, particularly trehalose synthesis, which can be a major node of Pi restoration (34, 37). We assessed the contribution of increased trehalose synthesis towards the high Pi in ubp3Δ, by estimating Pi levels in the absence of trehalose 6-phosphate phosphatase (Tps2), which catalyzes the Pi-releasing step in trehalose synthesis. Notably, loss of Tps2 in ubp3Δ decreased Pi (Figure 3E). There was no additive difference in Pi between tps2Δ and ubp3Δtps2Δ (Figure 3E). As an added control, we assessed trehalose in WT and ubp3Δ in the absence of Tps2, and found no difference (Figure 3-figure supplement 1B). Therefore, increasing G6P flux towards trehalose biosynthesis is a major source of the increased Pi in ubp3Δ.

Revised version: We next asked if the increased Pi in ubp3Δ is because of altered G6P allocations towards different end-points, particularly trehalose synthesis, which can be a major node of Pi restoration (Gupta et al., 2019; Van Heerden et al., 2014). We assessed the contribution of increased trehalose synthesis towards the high Pi in ubp3Δ, by estimating Pi levels in the absence of trehalose 6-phosphate phosphatase (Tps2), which catalyzes the Pi-releasing step in trehalose synthesis. Notably, loss of Tps2 in ubp3Δ decreased Pi (Figure 3E). There was no additive difference in Pi between tps2Δ and ubp3Δtps2Δ (Figure 3E). As an added control, we assessed trehalose in WT and ubp3Δ in the absence of Tps2, and found no difference (Figure 3-figure supplement 1B). Therefore, increasing G6P flux towards trehalose biosynthesis is a major source of the increased Pi in ubp3Δ.

Since the major GAPDH isozymes, Tdh2 and Tdh3, are reduced in ubp3Δ, we directly asked if reducing GAPDH can decrease Pi consumption and increase Pi. To assess this, we generated tdh2Δtdh3Δ cells, which exhibit a growth defect, but are viable, permitting further analysis. tdh2Δtdh3Δ had higher Pi in high glucose (Figure 3F). Expectedly, we observed a decrease in ethanol in tdh2Δtdh3Δ (Figure 3-figure supplement 1C), along with an accumulation of F1,6 BP and G3P, and decreased 3PG and PEP (Figure 3-figure supplement 1D). However, G6P and trehalose levels between WT and tdh2Δtdh3Δ were comparable (Figure 3-figure supplement 1D, E). These data sugest that unlike in ubp3Δ, the increased Pi in tdh2Δtdh3Δ comes mainly from decreased Pi consumption (GAPDH step). To further assess the role of reduced glycolytic flux in increasing Pi (ubp3Δ and tdh2Δtdh3Δ), we measured the Pi in these cells growing in a gluconeogenic medium – 2%, ethanol. In this scenario, the GAPDH catalyzed reaction will be reversed, converting 1,3BPG to G3P, which should release and not consume Pi. Compared to WT, the Pi levels decreased in ubp3Δ and tdh2Δtdh3Δ (Figure 3G), suggesting that the changes in Pi in these mutants is driven by the relative change in Pi release vs consumption.

These results collectively indicate that the combined effect of increased Pi release coming from trehalose synthesis, and decreased Pi consumption from reduced GAPDH increase Pi levels in ubp3Δ.

Mitochondrial Pi availability correlates with mitochondrial activity in ubp3Δ

Earlier version: We therefore wondered if this observed Pi increase from the combined rewiring of glucose metabolism resulted in more Pi becoming accessible to the mitochondria. This would effectively result in Pi budgeting between cytosolic glycolysis and mitochondria based on the relative flux in different arms of glucose metabolism, by increasing Pi availability for the mitochondria (Figure 4A). To test if this were so, we first asked if the mitochondria in ubp3Δ have increased Pi. Mitochondria were isolated by immunoprecipitation from WT and ubp3Δ, the isolation efficiency was analyzed (Figure 4-figure supplement 1A), and relative Pi levels compared. Pi levels were normalised to Idh1 (isocitrate dehydrogenase) in isolated mitochondria, since Idh1 protein levels did not decrease in ubp3Δ (Figure 4-figure supplement 1B), and consistent with (28). Mitochondrial Pi was higher in ubp3Δ (Figure 4B). We next asked if the increased mitochondrial activity in ubp3Δ is a consequence of high Pi. If this were so, tdh2Δtdh3Δ should partially phenocopy ubp3Δ with respect to mitochondrial activity. Consistently, tdh2Δtdh3Δ have higher mitotracker intensity as well as Cox2 levels compared to WT (Figure 4-figure supplement 1C, Figure 4C). Also consistent with this, a high basal OCR in tdh2Δtdh3Δ was observed (Figure 4D), together indicating high mitochondrial activity in tdh2Δtdh3Δ, which is comparable to ubp3Δ.

Revised version: We therefore wondered if this observed Pi increase from the combined rewiring of glucose metabolism resulted in more Pi becoming accessible to the mitochondria. This would effectively result in Pi budgeting between cytosolic glycolysis and mitochondria based on the relative flux in different arms of glucose metabolism, by increasing Pi availability for the mitochondria (Figure 4A). If this were indeed so, a prediction would be that the pool of Pi in the mitochondria of ubp3Δ cells would be higher than WT cells, while the opposite would be expected in the cytosolic fraction of these cells, due to increased Pi in the mitochondria. To test if this were so, we first measured the cytosolic pools of Pi in WT and ubp3Δ. The cytosolic fraction was isolated (extensive experimental details are in the SI appendix) from WT and ubp3Δ cells, and Pi levels in this fraction were estimated (Figure 4B, Figure 4-figure supplement 1B). We observed significantly reduced Pi in the cytosolic fraction in ubp3Δ cells. Since total cellular Pi amounts is higher in ubp3Δ cells (Figure 3), a decreased cytolosolic Pi would be consistent with greater transport of Pi from cytosol to other organelles such as vacuole (where it is stored as polyphosphate) and/or mitochondria. Therefore, we next asked if the mitochondria in ubp3Δ have correspondingly increased Pi. Mitochondria were isolated by immunoprecipitation from WT and ubp3Δ, the isolation efficiency was analyzed (Figure 4-figure supplement 1A), and relative Pi levels compared. Pi levels were normalised to Idh1 (isocitrate dehydrogenase) in isolated mitochondria, since Idh1 protein levels did not decrease in ubp3Δ (Figure 4-figure supplement 1B), and consistent with (Isasa et al., 2015). Mitochondrial Pi was higher in ubp3Δ (Figure 4B). We next asked if the increased mitochondrial activity in ubp3Δ is a consequence of high Pi. If this were so, tdh2Δtdh3Δ should partially phenocopy ubp3Δ with respect to mitochondrial activity. Consistently, tdh2Δtdh3Δ have higher mitotracker intensity as well as Cox2 levels compared to WT (Figure 4-figure supplement 1C, Figure 4C). Also consistent with this, a high basal OCR in tdh2Δtdh3Δ was observed (Figure 4D), together indicating high mitochondrial activity in tdh2Δtdh3Δ, which is comparable to ubp3Δ.

Earlier version: Mitochondrial Pi availability correlates with mitochondrial activity in ubp3Δ A) A hypothetical mechanism of cytosolic free Pi controling mitochondrial activity by regulating mitochondrial Pi availability.

Revised version: Mitochondrial Pi availability correlates with mitochondrial activity in ubp3Δ

Added in the revised version: A) A hypothetical mechanism of cytosolic free Pi controling mitochondrial activity by regulating mitochondrial Pi availability.

Earlier version: B) Mitochondrial Pi amounts in WT vs ubp3Δ. Mitochondria were isolated by immunoprecipitation from WT and ubp3Δ and mitochondrial Pi estimated. Mitochondrial Pi levels (normalised to Idh1) are shown. Data represent mean ± SD from three biological replicates (n=3). Also see figure 4-figure supplement 1A, B.

Revised version: B) Cytosolic and mitochondrial Pi amounts in WT vs ubp3Δ. The cytosolic fraction was isolated by centrifugation (see SI appendix), and in separate experiments, mitochondria were isolated by immunoprecipitation from WT and ubp3Δ and mitochondrial Pi estimated. (i) Cytosolic Pi levels (relative as well as absolute) and (ii) mitochondrial Pi levels (normalised to Idh1) are shown. Data represent mean ± SD from three biological replicates (n=3) respectively for the cytosolic and mitochondrial measurements. Also see Figure 4-figure supplement 1A, B.

C) Cox2 protein in tdh2Δtdh3Δ. WT, ubp3Δ, and tdh2Δtdh3Δ were grown in high glucose and Cox2 protein was estimated. A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD.

D) Basal OCR levels in tdh2Δtdh3Δ. WT, ubp3Δ, and tdh2Δtdh3Δ were grown in high glucose and basal OCR was measured from two independent experiments (n=2). Data represent mean ± SD. Also see Figure 4-figure supplement 1C.

E) Comparative Pi amounts and Cox2 levels in ubp3Δ, tdh2Δtdh3Δ, WT cells. WT cells were grown in high glucose, ubp3Δ and tdh2Δtdh3Δ were grown in high glucose and low Pi, and Cox2 protein was estimated. A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD. Also see figure Figure 4-figure supplement 1D, F.

F) Pi amounts and basal OCR in ubp3Δ and tdh2Δtdh3Δ vs WT cells. WT cells were grown in high glucose, ubp3Δ and tdh2Δtdh3Δ were grown in high glucose and low Pi, and basal OCR was measured from three independent experiments (n=3). Data represent mean ± SD.

G) Effect of loss of mitochondrial Pi transporter Mir1 on Cox2protein. WT, ubp3Δ, mir1Δ, and mir1Δubp3Δ were grown in high glucose and Cox2 amounts compared. A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD.

H) Relationship of mitochondrial Pi transport and basal OCR in WT vs ubp3Δ. WT, ubp3Δ, mir1Δ, and mir1Δubp3Δ cells were grown in high glucose and basal OCR was measured from two independent experiments (n=2). Data represent mean ± SD.

Data information: *p<0.05, **p<0.01, ****p<0.0001.

We next asked if higher Pi is necessary to increase mitochondrial activity in ubp3Δ and tdh2Δtdh3Δ. Since glycolysis is defective in ubp3Δ and tdh2Δtdh3Δ, it is necessary to distinguish the effect of high Pi vs. only the effect of low glycolysis in activating mitochondria. Logically, if decreased glycolysis (independent of Pi) is sufficient to activate mitochondria, bringing down the Pi levels in ubp3Δ to that of WT should not affect mitochondrial activity. Notably, ubp3Δ grown in low (1 mM) Pi have Pi levels similar to WT in standard (normal Pi) medium (Figure 4-figure supplement 1D). ubp3Δ grown in low Pi also had decreased ethanol, suggesting reduced glycolysis (Figure 4-figure supplement 1E). Therefore, we used this condition to further understand the role of Pi in inducing mitochondrial activity. Mitotracker fluorescence decreased in both ubp3Δ and tdh2Δtdh3Δ in low Pi (Figure 4-figure supplement 1F). Note: basal mitotracker fluorescence in WT also decreases in low Pi, which is consistent with a required role of Pi for mitochondrial activity (Figure 4-figure supplement 1F). Similarly, Cox2 levels were reduced in both ubp3Δ and tdh2Δtdh3Δ (Figure 4E). Consistent with both reduced mitotracker intensity and Cox2 levels, basal OCR also decreased in both ubp3Δ and tdh2Δtdh3Δ in low Pi (Figure 4F). As an additional control, we used Rho0 strains (which have no mitochondrial DNA and therefore lack functional ETC) to compare basal OCR. The expectation in these strains is that basal OCR will not change if Pi changes. Consistently, we did not observe any significant difference in basal OCR in WT, ubp3Δ and ubp3Δ in low Pi in a Rho0 strain background (which lacks mitochondrial DNA) (Figure 4-figure supplement 1G). These data collectively suggest that high intracellular Pi is necessary to increase mitochondrial activity in ubp3Δ and tdh2Δtdh3Δ.

Earlier version: Next, we asked how mitochondrial Pi transport regulates mitochondrial activity. Mir1 and Pic2 are mitochondrial Pi transporters, with Mir1 being the major Pi transporter (44, 45). We first limited mitochondrial Pi availability in ubp3Δ by knocking-out MIR1. In mir1Δ, the increased Cox2 observed in ubp3Δ was no longer observed (Figure 4G). Consistent with this, we observed no further increase in the basal OCR in mir1Δubp3Δ compared to mir1Δ (Figure 4H). Furthermore, mir1Δ showed decreased Cox2 as well as basal OCR even in WT cells (Figure 4G, Figure 4H). These data together suggest that mitochondrial Pi transport is critical for increasing mitochondrial activity in ubp3Δ, and in maintaining basal mitochondrial activity even in high glucose.

Revised version: Next, we asked how mitochondrial Pi transport regulates mitochondrial activity. Mir1 and Pic2 are mitochondrial Pi transporters, with Mir1 being the major Pi transporter (Murakami et al., 1990; Zara et al., 1996). We first limited mitochondrial Pi availability in ubp3Δ by knocking-out MIR1. In mir1Δ, the increased Cox2 observed in ubp3Δ was no longer observed (Figure 4G). Consistent with this, we observed no further increase in the basal OCR in mir1Δubp3Δ compared to mir1Δ (Figure 4H). Furthermore, mir1Δ showed decreased Cox2 as well as basal OCR even in WT cells (Figure 4G, Figure 4H). These data together suggest that mitochondrial Pi transport is critical for increasing mitochondrial activity in ubp3Δ, and in maintaining basal mitochondrial activity even in high glucose.

As a control, no significant increase in Mir1 and Pic2 was observed in ubp3Δ (Figure 4-figure supplement 1H), suggesting that ubp3Δ do not increase mitochondrial Pi by merely increasing the Pi transporters, but rather by increasing available Pi pools.

Earlier version: Taken together, these data suggest the possibility that increased cytosolic Pi increases the mitochondrial Pi pool, which increases mitochondrial activity. Decreasing mitochondrial Pi by either reducing total Pi, or by reducing mitochondrial Pi transport decreases mitochondrial activity.

Revised version: Taken together, these data suggest the possibility that the altered Pi homeostais in ubp3Δ cells increases the mitochondrial Pi pool. This increased mitochondrial Pi pool correlates with increased mitochondrial activity. Decreasing mitochondrial Pi by either reducing total Pi, or by reducing mitochondrial Pi transport decreases mitochondrial activity.

Mitochondrial Pi availability constrains mitochondrial activity under high glucose

So far, these data suggest that the cytosolic Pi available for the mitochondria can determine the extent of mitochondrial activity. Therefore, we further investigated if mitochondrial Pi allocation was a necessary constraint for glucose-mediated mitochondrial repression.

Earlier version: To test this, we first asked how important mitochondrial Pi transport was to switch to increased respiration. In WT yeast, low glucose or glycolytic inhibition will result in increased respiration (22). What happens therefore if we restrict mitochondrial Pi in this context? For this, we measured the basal OCR in WT and mir1Δ after switching from high (2%) to low (0.1%) glucose. We observed a significant increase in the basal OCR in WT but not in mir1Δ (Figure 5A). The alternate scenario is after glycolytic inhibition. We assessed the role of mitochondrial Pi in this context, by inhibiting glycolysis using 2-deoxyglucose (2DG). WT, but not mir1Δ, increased their OCR (respiration) upon a 1-hour treatment with 2DG (Figure 5B). Consistent with this, mitotracker fluorescence increased with an increase in 2DG in WT, but not in mir1Δ (Figure 5-figure supplement 1A). We further asked if the mitochondrial Pi transporter itself glucose repressed, and therefore assessed Mir1 amounts in high and low glucose. Mir1 levels are higher upon a shift to low glucose, and in cells grown in 2% ethanol, suggesting that Mir1 is glucose repressed (Figure 5C, Figure 5-figure supplement 1B). These data suggest that mitochondrial Pi transport is necessary for increasing mitochondrial activity after glucose derepression.

Revised version: To test this, we first asked how important mitochondrial Pi transport was to switch to increased respiration. In WT yeast, low glucose or glycolytic inhibition will result in increased respiration (Broach, 2012). What happens therefore if we restrict mitochondrial Pi in this context? For this, we measured the basal OCR in WT and mir1Δ after switching from high (2%) to low (0.1%) glucose. We observed a significant increase in the basal OCR in WT but not in mir1Δ (Figure 5A). The alternate scenario is after glycolytic inhibition. We assessed the role of mitochondrial Pi in this context, by inhibiting glycolysis using 2-deoxyglucose (2DG). WT, but not mir1Δ, increased their OCR (respiration) upon a 1-hour treatment with 2DG (Figure 5B). Consistent with this, mitotracker fluorescence increased with an increase in 2DG in WT, but not in mir1Δ (Figure 5-figure supplement 1A). We further asked if the mitochondrial Pi transporter itself glucose repressed, and therefore assessed Mir1 amounts in high and low glucose. Mir1 levels are higher upon a shift to low glucose, and in cells grown in 2% ethanol, suggesting that Mir1 is glucose repressed (Figure 5C, Figure 5-figure supplement 1B). These data suggest that mitochondrial Pi transport is necessary for increasing mitochondrial activity after glucose derepression.

Mitochondrial Pi availability constrains mitochondrial activity under high glucose

A) Relationship of mitochondrial Pi transport and respiration after glucose removal. WT and mir1Δ cells were cultured in high (2%) glucose and shifted to low (0.1%) glucose for 1 hour. The normalized basal OCR, from two independent experiments (n=2) are shown. Data represent mean ± SD.

B) Requirement of mitochondrial Pi transport for switch to respiration upon glycolytic inhibition by 2DG. WT and mir1Δ cells were cultured in high glucose and treated with or without 0.25% 2DG for 1 hour. Basal OCR was measured from two independent experiments (n=2). Data represent mean ± SD. Also see Figure 5-figure supplement 1A.

C) Glucose-dependent regulation of Mir1. Cells (with Mir1-HA) were grown in high glucose and shifted to low glucose (0.1% glucose) for 1 hour, and Mir1 levels compared. A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD. Also see Figure 5-figure supplement 1B.

D) Increasing Pi concentrations and mitochondrial activity in isolated mitochondria. Mitochondria were isolated from WT cells grown in high glucose, incubated with 1 mM pyruvate, 1 mM malate, 0.5 mM ADP and 0-50 mM KH2PO4. The mitochondrial activity was estimated by mitotracker fluorescence intensity, and intensities relative to the sample with 0 mM KH2PO4 is shown. Data represent mean ± SD from three biological replicates (n=3).

E) Effect of overexpressing Mir1 on Cox2 protein. WT (containing empty vector) and Mir1 overexpressing (Mir1OE) cells were grown in high glucose and Cox2 levels were estimated. A representative blot (out of three biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD. Also see figure S5F.

F) Effect of overexpressing Mir1 on basal OCR. The basal OCR in WT (containing empty vector) and Mir1OE in high glucose was measured from three independent experiments (n=3). Data represent mean ± SD.

Data information: *p<0.05, **p<0.01, ***p<0.001.

Earlier version: We next asked whether just adding external Pi was sufficient to increase mitochondrial activity, when cells are in high glucose. In medium supplemented with excess Pi, the internal Pi increases as seen earlier (Figure 3C). Therefore, a simplistic assumption would be that the addition of external Pi to cells in high glucose would also increase mitochondrial Pi. However, an alternate possibility presents itself whererin since glycolytic flux is already high in glucose, supplementing Pi will continue to fuel glycolysis. Indeed, this was originally obseved by Harden and Young in 1908, where adding Pi increased fermentation (46). In such a scenario, mitochondrial Pi will not increase. We estimated mitochondrial Pi in this condition where excess Pi was externally supplemented. Notably, cells grown in high Pi had decreased mitochondrial Pi (Figure 5-figure supplement 1C), without any changes in total mitochondria volume or amounts (Figure 5-figure supplement 1D). Furthermore, directly adding Pi to cells growing in high glucose also decreased basal OCR (Figure 5-figure supplement 1E), consistent with decreased mitochondrial Pi. These data indicate that in high glucose, simply supplementing Pi will not increase Pi access to the mitochondria. We therefore now asked, if we inhibit glycolytic flux and then supplement Pi, what would happen to mitochondrial activity. For this, we treated cells with 2DG, and subsequently added Pi and measured the OCR (Figure 5-figure supplement 1F). In this case, supplementing Pi increased the basal OCR (Figure 5-figure supplement 1F). Collectively, these data suggest that a combination of decreasing glycolysis and increasing Pi can together increase respiration. Next, in order to directly test mitochondrial activation based on external Pi availability, we isolated mitochondria, and estimated activity in vitro upon adding increasing Pi. Mitochondrial activity increased with increased Pi, with maximum activity observed with 25 mM Pi supplemented (Figure 5D). In a complementary experiment, we overexpressed the Mir1 transporter in wild-type cells, to increase Pi within mitochondria (Figure 5-figure supplement 1G). Mir1-OE cells have higher Cox2 levels and basal OCR (Figure 5E, 5F). Therefore, increasing Pi transport to mitochondria is sufficient to increase mitochondrial activity in high glucose.

Revised version: We next asked whether just adding external Pi was sufficient to increase mitochondrial activity, when cells are in high glucose. In medium supplemented with excess Pi, the internal Pi increases as seen earlier (Figure 3C). Therefore, a simplistic assumption would be that the addition of external Pi to cells in high glucose would also increase mitochondrial Pi. However, an alternate possibility presents itself whererin since glycolytic flux is already high in glucose, supplementing Pi will continue to fuel glycolysis. Indeed, this was originally obseved by Harden and Young in 1908, where adding Pi increased fermentation (Harden and Young, 1906). In such a scenario, there could be an increase in the cytosolic Pi but not the mitochondrial Pi. We estimated the cytosolic and mitochondrial Pi in this condition where excess Pi was externally supplemented. Notably, cells grown in high Pi had increased cytosolic Pi, but decreased mitochondrial Pi (Figure 5-figure supplement 1C), without any changes in total mitochondria volume or amounts (Figure 5-figure supplement 1D). Furthermore, directly adding Pi to cells growing in high glucose also decreased basal OCR (Figure 5-figure supplement 1E), consistent with decreased mitochondrial Pi. These data indicate that in high glucose, simply supplementing Pi will not increase Pi access to the mitochondria, and instead results in an accumulation of Pi in the cytosol (Figure 5-figure supplement 1C). We therefore now asked, if we inhibit glycolytic flux and then supplement Pi, what would happen to mitochondrial activity. For this, we treated cells with 2DG, and subsequently added Pi and measured the OCR (Figure 5-figure supplement 1F). In this case, supplementing Pi increased the basal OCR (Figure 5-figure supplement 1F). Collectively, these data suggest that a combination of decreasing glycolysis and increasing Pi can together increase respiration. Next, in order to directly test mitochondrial activation based on external Pi availability, we isolated mitochondria, and estimated activity in vitro upon adding increasing Pi. Mitochondrial activity increased with increased Pi, with maximum activity observed with 25 mM Pi supplemented (Figure 5D). In a complementary experiment, we overexpressed the Mir1 transporter in wild-type cells, to increase Pi within mitochondria (Figure 5-figure supplement 1G). Mir1-OE cells have higher Cox2 levels and basal OCR (Figure 5E, 5F). Therefore, increasing Pi transport to mitochondria is sufficient to increase mitochondrial activity in high glucose.

Earlier version: Mitochondrial pyruvate transport is also required for mitochondrial respiration (47). We asked where Pi availability stands in a hierarchy of constraints for mitochondrial derepression, as compared to mitochondrial pyruvate transport. We measured the amounts of the Mpc3 subunit of the mitochondrial pyruvate carrier (MPC) complex (47, 48). Mpc3 protein increases in ubp3Δ (Figure 5-figure supplement 1H). This also correlated with the unimpaired TCA cycle flux (Figure 2-figure supplement 2D) and the increased mitochondrial activity. Interestingly, in ubp3Δ grown in low Pi, Mpc3 further increased (Figure 5-figure supplement 1H), but as shown earlier this condition cannot increase OCR or mitochondrial activity (Figure 4F, Figure 4-figure supplement 1F). Basal Mpc3 levels decrease in mir1Δ, but upon shifting to 0.1% glucose, Mpc3 increases in both WT and mir1Δ, with higher levels in mir1Δ (Figure 5-figure supplement 1I). Therefore, even where Mpc3 is high (ubp3Δ in low Pi, and mir1Δ in low glucose), mitochondrial activity remains low if Pi is restricted (Figure 4F, 4H). There was also no decrease in basal OCR in mpc3Δ in high glucose, and the basal OCR increased to the same level as of WT after shifting to 0.1% glucose (Figure 5-figure supplement 1J). Since Mpc3 changes with mitochondrial Pi availability (Figure S5G, S5H), we also measured Mpc3 in the Mir1OE. No further changes in Mpc3 were observed in Mir1OE (Figure 5-figure supplement 1K), indicating that increasing mitochondrial Pi alone need not increase Mpc3. Overall, although Mpc3 levels corelate with decreased glycolysis (ubp3Δ - Figure 5-figure supplement 1H, ubp3Δ in low Pi- Figure 5-figure supplement 1H, low glucose - Figure 5-figure supplement 1I), increased Mpc3 alone cannot increase mitochondrial activity and respiration in the absence of adequate mitochondrial Pi.

Revised version: Mitochondrial pyruvate transport is also required for mitochondrial respiration (Timoń-Gómez et al., 2013). We asked where Pi availability stands in a hierarchy of constraints for mitochondrial derepression, as compared to mitochondrial pyruvate transport. We measured the amounts of the Mpc3 subunit of the mitochondrial pyruvate carrier (MPC) complex (Bender et al., 2015; Timoń-Gómez et al., 2013). Mpc3 protein increases in ubp3Δ (Figure 5-figure supplement 1H). This also correlated with the unimpaired TCA cycle flux (Figure 2-figure supplement 2D) and the increased mitochondrial activity. Interestingly, in ubp3Δ grown in low Pi, Mpc3 further increased (Figure 5-figure supplement 1H), but as shown earlier this condition cannot increase OCR or mitochondrial activity (Figure 4F, Figure 4-figure supplement 1F). Basal Mpc3 levels decrease in mir1Δ, but upon shifting to 0.1% glucose, Mpc3 increases in both WT and mir1Δ, with higher levels in mir1Δ (Figure 5-figure supplement 1I). Therefore, even where Mpc3 is high (ubp3Δ in low Pi, and mir1Δ in low glucose), mitochondrial activity remains low if Pi is restricted (Figure 4F, 4H). There was also no decrease in basal OCR in mpc3Δ in high glucose, and the basal OCR increased to the same level as of WT after shifting to 0.1% glucose (Figure 5-figure supplement 1J). Since Mpc3 changes with mitochondrial Pi availability (Figure S5G, S5H), we also measured Mpc3 in the Mir1OE. No further changes in Mpc3 were observed in Mir1OE (Figure 5-figure supplement 1K), indicating that increasing mitochondrial Pi alone need not increase Mpc3. Overall, although Mpc3 levels corelate with decreased glycolysis (ubp3Δ - Figure 5-figure supplement 1H, ubp3Δ in low Pi-Figure 5-figure supplement 1H, low glucose - Figure 5-figure supplement 1I), increased Mpc3 alone cannot increase mitochondrial activity and respiration in the absence of adequate mitochondrial Pi.

Collectively, mitochondrial Pi availability constrains glucose-mediated mitochondrial repression. Increasing available pools of Pi to enter the mitochondria is sufficient to induce mitochondrial activity.

Repression of mitochondrial respiration via Pi budgeting is conserved in Ubp3 mutants across diverse yeast genetic backgrounds

Earlier version: So far, we have identified a role for intracellular Pi budgeting as a constraint for mitochondrial activity under high glucose. These were all carried out using a robust, prototrophic yeast strain from a CEN.PK background. S. cerevisiae however, while Crabtree positive, have tremendous genetic diversity (49). We therefore asked if this mitochondrial repression through Pi budgeting (mediated by Ubp3 function) is conserved across other strains of S. cerevsiae as well. To test this, we generated Ubp3 deletion mutants in different genetic backgrounds of S. cerevsiae, including BY4742, W303 and Σ1278. In all these strain backgrounds, we observed a significant increase in mitotracker fluorescne intensity and Cox2 protein levels in ubp3Δ (Figure 6A, B). This suggests the role of Ubp3 as a regulator of mitochondrial repression, independent of the genetic background of the yeast (S. cerevsiae) strain. To further assess if loss of Ubp3 shows a concurrent increase in Pi levels, we measured the total Pi levels in WT and ubp3Δ in a W303 strain background. We observed a significant increase in total Pi levels in ubp3Δ cells in this strain (Figure 6C), similar to what we observed in the CEN.PK strain (Figure 3C). Finally, to test if the altered Pi budgeting regulates the mitochondrial activity in these cells, we measured the basal OCR in WT, ubp3Δ and ubp3Δ in low (1 mM) Pi medium, in the W303 strain background. Consistent with the increase in mitotracker fluorescence intensity and Cox2 protein levels, we observed a significant increase in the basal OCR in ubp3Δ cells (Figure 6D). This increase was not observed in ubp3Δ in a low Pi medium (Figure 6D). This suggests that the role of altered Pi budgeting in regulating mitochondrial respiration is conserved in other genetic backgrounds of S. cerevisiae.

Revised version: So far, we have identified a role for intracellular Pi budgeting as a constraint for mitochondrial activity under high glucose. These were all carried out using a robust, prototrophic yeast strain from a CEN.PK background. S. cerevisiae however, while Crabtree positive, have tremendous genetic diversity (Peter et al., 2018). We therefore asked if this mitochondrial repression through Pi budgeting (mediated by Ubp3 function) is conserved across other strains of S. cerevsiae as well. To test this, we generated Ubp3 deletion mutants in different genetic backgrounds of S. cerevsiae, including BY4742, W303 and Σ1278. In all these strain backgrounds, we observed a significant increase in mitotracker fluorescne intensity and Cox2 protein levels in ubp3Δ (Figure 6A, B). This suggests the role of Ubp3 as a regulator of mitochondrial repression, independent of the genetic background of the yeast (S. cerevsiae) strain. To further assess if loss of Ubp3 shows a concurrent increase in Pi levels, we measured the total Pi levels in WT and ubp3Δ in a W303 strain background. We observed a significant increase in total Pi levels in ubp3Δ cells in this strain (Figure 6C), similar to what we observed in the CEN.PK strain (Figure 3C). Finally, to test if the altered Pi budgeting regulates the mitochondrial activity in these cells, we measured the basal OCR in WT, ubp3Δ and ubp3Δ in low (1 mM) Pi medium, in the W303 strain background. Consistent with the increase in mitotracker fluorescence intensity and Cox2 protein levels, we observed a significant increase in the basal OCR in ubp3Δ cells (Figure 6D). This increase was not observed in ubp3Δ in a low Pi medium (Figure 6D). This suggests that the role of altered Pi budgeting in regulating mitochondrial respiration is conserved in other genetic backgrounds of S. cerevisiae.

Repression of mitochondrial respiration via Pi budgeting is conserved in Ubp3 mutants across diverse yeast genetic backgrounds

A) Effect of loss of Ubp3 on mitochondrial membrane potential in different yeast strains. WT and ubp3Δ cells (in CEN.PK, BY4742, W303 and Σ1278 strains of S. cerevisiae) were grown in high glucose and relative mitochondrial membrane potential was measured. Data represent mean ± SD from three biological replicates (n=3).

B) Effect of loss of Ubp3 on ETC complex IV subunit Cox2. WT and ubp3Δ (in BY4742, W303 and Σ1278 strains of S. cerevisiae) were grown in high glucose, and Cox2 was measured. A representative blot (out of 3 biological replicates, n=3) and their quantifications are shown. Data represent mean ± SD.

C) Intracellular Pi levels in WT and ubp3Δ in W303 strain background. WT and ubp3Δ (in W303 strain background) were grown in high glucose and the total free phosphate (Pi) levels were estimated. Data represent mean ± SD from three biological replicates (n=3).

D) Effect of low Pi on the basal OCR in WT and ubp3Δ cells in W303 strain background. WT cells were grown in high glucose and ubp3Δ were grown in high glucose and low Pi, and basal OCR was measured. Data represent mean ± SD (n=3).

E) Requirement of mitochondrial Pi transport for growth after 2DG treatment. Shown are serial dilution growth assays in high glucose in the presence and absence of 0.1% 2DG, using WT and mir1Δ cells. The results after 40hrs incubation/30°C are shown.

F) A model illustrating how mitochondrial Pi availability controls mitochondrial activity. In high glucose, the decreased Pi due to high Pi consumption in glycolysis, along with the glucose-mediated repression of mitochondrial Pi transporters, decreases mitochondrial Pi availability. This reduces mitochondrial activity. In low glucose, increased mitochondrial Pi transporters and lower glycolytic flux increases mitochondrial Pi, leading to enhanced mitochondrial activity. In ubp3Δ cells in high glucose, high trehalose synthesis and lower glycolytic flux results in an increase in Pi. This increases mitochondrial Pi availability and thereby the mitochondrial activity.

Data information: *p<0.05, **p<0.01, ***p<0.001.

Finally, we asked how important mitochondrial Pi transport was for growth, under glycolytic inhibition. Consistent with the requirement for mitochondrial Pi transport to increase mitochondrial respiration upon glycolytic inhibition (Figure 5B), WT cells only exhibit a slightly decreased growth in the presence of 2DG (Figure 6E). In contrast, mir1Δ show a severe growth defect upon 2DG treatment (Figure 6E), revealing a synergetic effect of combining 2DG with inhibiting mitochondrial Pi transport. Therefore, the combined inhibition of glycolysis and mitochondrial Pi transport restricts the growth of glycolytic cells.

Collectively, our data suggests a conserved role for intracellular Pi bugeting in regulating mitochondrial repression in high glucose and the role of mitochondrial Pi transport in regulating adaptation for growth under glycolytic inhibition.

Discussion

Earlier version: In this study, we highlight a role for Pi budgeting between cytosolic glycolysis and mitochondrial processes (which compete for Pi) in constraining mitochondrial repression (Figure 6F). Ubp3 controls this process (Figure 1), by maintaining the amounts of the glycolytic enzymes Pfk1 and GAPDH (Tdh2 and Tdh3) and thereby allows high glycolytic flux. At high fermentation rates, glycolytic enzymes levels at maximal activity maintain high glycolytic rates (effectively following zero-order kinetics), and therefore, changes in the enzyme levels will have a direct, proportionate effect on flux (50). The loss of Ubp3 decreases glycolytic flux, resulting in a systems-level, mass-action based rewiring of glucose metabolism where more G6P is routed towards trehalose synthesis and PPP (Figure 2, Figure 6F). Indeed, inhibiting just phosphofructokinase can reroute glucose flux from glycolysis to PPP (51–53), and our study now permits contextualized interpretations of these results. Such reallocations of glucose flux will collectively increase overall Pi, coming from the combined effect of increased Pi release from trehalose synthesis, and decreased Pi consumption via reduced GAPDH. This altered intracellular Pi economy increases Pi pools available to mitochondria, and increasing mitochondrial Pi is necessary and sufficient to increase respiration in high glucose (Figure 4, Figure 5, Figure 6). Mitochondrial Pi transport maintains mitochondrial activity in high glucose, and increases mitochondrial activity in low glucose (Figure 4, Figure 5). Finally, the mitochondrial Pi transporter Mir1 itself decreases in high glucose (Figure 5). Therefore, this glucose-dependent repression of Mir1 also restricts mitochondrial Pi availability (and thereby activity) in high glucose.

Revised version: In this study, we highlight a role for Pi budgeting between cytosolic glycolysis and mitochondrial processes (which compete for Pi) in constraining mitochondrial repression (Figure 6F). Ubp3 controls this process (Figure 1), by maintaining the amounts of the glycolytic enzymes Pfk1 and GAPDH (Tdh2 and Tdh3) and thereby allows high glycolytic flux. At high fermentation rates, glycolytic enzymes levels at maximal activity maintain high glycolytic rates (effectively following zero-order kinetics), and therefore, changes in the enzyme levels will have a direct, proportionate effect on flux (Grigaitis and Teusink, 2022). The loss of Ubp3 decreases glycolytic flux, resulting in a systems-level, mass-action based rewiring of glucose metabolism where more G6P is routed towards trehalose synthesis and PPP (Figure 2, Figure 6F). Indeed, inhibiting just phosphofructokinase can reroute glucose flux from glycolysis to PPP (Hollinshead et al., 2016; Miyazawa et al., 2017; Yi et al., 2012), and our study now permits contextualized interpretations of these results. Such reallocations of glucose flux will collectively increase overall Pi, coming from the combined effect of increased Pi release from trehalose synthesis, and decreased Pi consumption via reduced GAPDH. This altered intracellular Pi economy increases Pi pools available to mitochondria, and increasing mitochondrial Pi is necessary and sufficient to increase respiration in high glucose (Figure 4, Figure 5, Figure 6). Mitochondrial Pi transport maintains mitochondrial activity in high glucose, and increases mitochondrial activity in low glucose (Figure 4, Figure 5). Finally, the mitochondrial Pi transporter Mir1 itself decreases in high glucose (Figure 5). Therefore, this glucose-dependent repression of Mir1 also restricts mitochondrial Pi availability (and thereby activity) in high glucose.

Traditionally, loss-of-function mutants of metabolic enzymes are used to understand metabolic state regulation. This approach negates nuanced investigations, since metabolic enzymes are often essential for viability. Furthermore, metabolic pathways have multiple, contextually regulated nodes, through which cells maintain their metabolic state. Therefore, alternate approaches to identify global regulators of metabolic states (as opposed to single enzymes), might uncover ways via which multiple nodes are simultaneously tuned, and can reveal unanticipated systems-level principles of metabolic state rewiring. In this study of glucose-mediated mitochondrial repression, the loss of the DUB Ubp3 decreases glycolytic flux by reducing the enzymes at two critical nodes in the pathway - Pfk1 and GAPDH. Unlike loss of function mutants, a reduction in amounts will only rewire metabolic flux. By ‘hitting’ multiple steps in glycolysis simultaneously, ubp3Δ have decreased Pi consumption, as well as increased Pi release. Such a cumulative phenomenon reveals more than inhibiting only GAPDH, where increased Pi comes only from reduced Pi consumption, and not from increased trehalose biosynthesis. Our serendipitous identification of a regulator which regulates multiple steps in glucose metabolism to change the metabolic environment, now suggests a general basis of mitochondrial regulation that would have otherwise remained hidden. Separately, finding the substrates of Ubp3 and whether Ubp3 directly regulates glycolytic enzymes are exciting future research questions requiring concurrent innovations in accessible chemical-biological approaches to study DUBs.

Earlier version: Because phosphates are ubiquitous, it is challenging to identify hierarchies of Pi-dependent processes in metabolic state regulation (35). Phosphate transfer reactions are the foundation of metabolism, driving multiple, thermodynamically unfavorable reactions (54, 55). Contextually, the laws of mass action predict that the relative rates of these reactions will regulate overall Pi balance, and contrarily the Pi allocation to Pi-dependent reactions will determine reaction rates (35, 37). Additionally, cells might control Pi allocations for different reactions via compartmentalizing Pi in organelles, to spatially restrict Pi availability (56–58). Our data collectively suggest a paradigm where the combination of factors regulates mitochondrial Pi and thereby activity. In glycolytic yeast cells growing in high glucose, mitochondrial Pi availability becomes restricted due to higher utilization of Pi in glycolysis compared to mitochondria. Consistent with this, a rapid decrease in Pi upon glucose addition has been observed (18, 19, 39). Interesting, in vitro studies with isolated mitochondria from tumor cells also find that decreasing Pi levels decreases respiration (19), which would be consistent with this scenario. Further, supplementing Pi correlates with decreased mitochondrial repression in tumors (18, 59). By increasing Pi through a systems-level rewiring of glucose metabolism (such as in ubp3Δ cells), cells can collectively increase mitochondrial access to Pi. This Pi budgeting determines mitochondrial activity. Supplementing Pi under conditions of low glycolysis (where mitochondrial Pi transport is enhanced), as well as directly supplementing Pi to isolated mitochondria, increases respiration (Figure 5, Figure 5-figure supplement 1). Therefore, in order to derepress mitochondria, a combination of increased Pi along with decreased glycolysis is required. An additional systems-level phenomenon that might regulate Pi transport to the mitochondria is the decrease in cytosolic pH upon decreased glycolysis (60, 61). The cytosolic pH in highly glycolytic cells is ∼7, and decreasing glycolysis results in cytosolic acidification (60, 61). Therefore, under conditions of decreased glycolysis (2DG treatment, deletion of Ubp3, and decreased GAPDH activity), cytosolic pH becomes acidic. Since mitochondrial Pi transport itself is dependent on the proton gradient, a low cytosolic pH would favour mitochondrial Pi transport (62). Therefore, under conditions of decreased glycolysis (2DG treatment, or loss of Ubp3, or decreased GAPDH activity), where cytosolic pH would be acidic, increasing cytosolic Pi might indirectly increase mitochondria Pi transport, thereby leading to increased respiration. Alternately, increasing mitochondrial Pi transporter amounts can achieve the same result, as seen by overexpressing Mir1 (Figure 5). A similar observation has been reported in Arabidopsis, reiterating an evolutionarily conserved role for mitochondrial Pi in controlling respiration (63). Relatedly, glycolytic inhibition can suppress cell proliferation in Warburg-positive tumors (64, 65), or inflammatory responses (66). However, these cells survive by switching to mitochondrial respiration (67, 68), requiring alternate approaches to prevent their proliferation (69). Inhibiting mitochondrial Pi transport in combination with glycolytic inhibition could restrict the proliferation of Warburg/Crabtree positive cells.

Revised version: Because phosphates are ubiquitous, it is challenging to identify hierarchies of Pi-dependent processes in metabolic state regulation (Gupta and Laxman, 2021). Phosphate transfer reactions are the foundation of metabolism, driving multiple, thermodynamically unfavorable reactions (Kamerlin et al., 2013; Westheimer, 1987). Contextually, the laws of mass action predict that the relative rates of these reactions will regulate overall Pi balance, and contrarily the Pi allocation to Pi-dependent reactions will determine reaction rates (Gupta and Laxman, 2021; Van Heerden et al., 2014). Additionally, cells might control Pi allocations for different reactions via compartmentalizing Pi in organelles, to spatially restrict Pi availability (Booth and Guidotti, 1997; Solesio et al., 2021; Vila et al., 2022). Our data collectively suggest a paradigm where the combination of factors regulates mitochondrial Pi and thereby activity. In glycolytic yeast cells growing in high glucose, mitochondrial Pi availability becomes restricted due to higher utilization of Pi in glycolysis compared to mitochondria. Consistent with this, a rapid decrease in Pi upon glucose addition has been observed (Hohmann et al., 1996; Koobs, 1972; Rodríguez-Enríquez et al., 2001). Interesting, in vitro studies with isolated mitochondria from tumor cells also find that decreasing Pi levels decreases respiration (Rodríguez-Enríquez et al., 2001), which would be consistent with this scenario. Further, supplementing Pi correlates with decreased mitochondrial repression in tumors (Brin and McKee, 1956; Koobs, 1972). By increasing Pi through a systems-level rewiring of glucose metabolism (such as in ubp3Δ cells), cells can collectively increase mitochondrial access to Pi. This Pi budgeting determines mitochondrial activity. Supplementing Pi under conditions of low glycolysis (where mitochondrial Pi transport is enhanced), as well as directly supplementing Pi to isolated mitochondria, increases respiration (Figure 5, Figure 5-figure supplement 1). Notably, this increased respiration does not happen upon directly supplementing Pi to highly glycolytic WT cells, where the Pi increases in cytosol, without increasing mitochondrial Pi (Figure 5-figure supplement 1C). Therefore, in order to derepress mitochondria, a combination of increased Pi along with decreased glycolysis is required. An additional systems-level phenomenon that might regulate Pi transport to the mitochondria is the decrease in cytosolic pH upon decreased glycolysis (Dechant et al., 2010; Orij et al., 2011). The cytosolic pH in highly glycolytic cells is ∼7, and decreasing glycolysis results in cytosolic acidification (Dechant et al., 2010; Orij et al., 2011). Therefore, under conditions of decreased glycolysis (2DG treatment, deletion of Ubp3, and decreased GAPDH activity), cytosolic pH becomes acidic. Since mitochondrial Pi transport itself is dependent on the proton gradient, a low cytosolic pH would favour mitochondrial Pi transport (Hamel et al., 2004). Therefore, under conditions of decreased glycolysis (2DG treatment, or loss of Ubp3, or decreased GAPDH activity), where cytosolic pH would be acidic, increasing cytosolic Pi might indirectly increase mitochondria Pi transport, thereby leading to increased respiration. Alternately, increasing mitochondrial Pi transporter amounts can achieve the same result, as seen by overexpressing Mir1 (Figure 5). A similar observation has been reported in Arabidopsis, reiterating an evolutionarily conserved role for mitochondrial Pi in controlling respiration (Jia et al., 2015). Relatedly, glycolytic inhibition can suppress cell proliferation in Warburg-positive tumors (O’Neill et al., 2019; Pelicano et al., 2006), or inflammatory responses (Soto-Heredero et al., 2020). However, these cells survive by switching to mitochondrial respiration (Lu et al., 2015; Shiratori et al., 2019), requiring alternate approaches to prevent their proliferation (Cheng et al., 2012). Inhibiting mitochondrial Pi transport in combination with glycolytic inhibition could restrict the proliferation of Warburg/Crabtree positive cells. It is important to highlight that our experiments, whether involving Pi supplementation or Pi limitations, maintain the cellular Pi concentration within the millimolar range, and are conducted within a short timeframe (∼ 1 hour). This differs significantly from Pi starvation studies, where cells are subjected to prolonged and complete Pi deprivation. In those contexts, cells trigger extensive metabolic adaptations in order to sustain available Pi pools including an increase in mitochondrial membrane potential which can be independent of respiration (Ouyang et al., 2024).

Earlier version: Since its discovery in the 1920s, the phenomenon of accelerated glycolysis with concurrent mitochondrial repression has been intensely researched. Yet, the biochemical constraints for glucose-mediated mitochondrial repression remains unresolved. One hypothesis suggests that the availability of glycolytic intermediates might determine the extent of mitochondrial repression. F1,6BP inhibits complex III and IV of the electron transport chain (ETC) in Crabtree positive yeast (14–17). Similarly, the ratio between G6P and F1,6BP regulates the extent of mitochondrial repression (16, 17). Although G6P/F6P accumulates in ubp3Δ (Figure 2C), this is not the case in tdh2Δtdh3Δ (GAPDH mutant) (Figure 3-figure supplement 1D), suggesting that G6P/F6P accumulation in itself is not the criterion to increase mitochondrial activity. Separately, the competition for common metabolites/co-factors between glycolysis and respiration (such as ADP, Pi or pyruvate) could drive this phenomenon (14, 18). Here we observe that Mpc3 mediated mitochondrial pyruvate transport alone cannot increase respiration. An additional consideration is the possible contribution of changes in ADP in regulating mitochondrial activity, where the use of ADP in glycolysis might limit mitochondrial ADP. Therefore, when Pi changes as a consequence of glycolysis, it coud be imagined that a change in ADP balance can coincidentally occur. However, prior studies show that even though cytosolic ADP decreases in the presence of glucose, this does not limit mitochondrial ADP uptake, or decrease respiration, due to the very high affinity of the mitochondrial ADP transporter (14, 19). These collectively reiterate the importance of Pi access and transport to mitochondria in constraining mitochondrial respiration. Indeed, this interpretation can also contextually explain observations from other model systems where mitochondrial Pi transport seems to regulate respiration (70, 71).

Revised version: Since its discovery in the 1920s, the phenomenon of accelerated glycolysis with concurrent mitochondrial repression has been intensely researched. Yet, the biochemical constraints for glucose-mediated mitochondrial repression remains unresolved. One hypothesis suggests that the availability of glycolytic intermediates might determine the extent of mitochondrial repression. F1,6BP inhibits complex III and IV of the electron transport chain (ETC) in Crabtree positive yeast (Diaz-Ruiz et al., 2011; Díaz-Ruiz et al., 2008; Hammad et al., 2016; Lemus et al., 2018). Similarly, the ratio between G6P and F1,6BP regulates the extent of mitochondrial repression (Díaz-Ruiz et al., 2008; Lemus et al., 2018). Although G6P/F6P accumulates in ubp3Δ (Figure 2C), this is not the case in tdh2Δtdh3Δ (GAPDH mutant) (Figure 3-figure supplement 1D), suggesting that G6P/F6P accumulation in itself is not the criterion to increase mitochondrial activity. Separately, the competition for common metabolites/co-factors between glycolysis and respiration (such as ADP, Pi or pyruvate) could drive this phenomenon (Diaz-Ruiz et al., 2011; Koobs, 1972). Here we observe that Mpc3 mediated mitochondrial pyruvate transport alone cannot increase respiration. An additional consideration is the possible contribution of changes in ADP in regulating mitochondrial activity, where the use of ADP in glycolysis might limit mitochondrial ADP. Therefore, when Pi changes as a consequence of glycolysis, it coud be imagined that a change in ADP balance can coincidentally occur. However, prior studies show that even though cytosolic ADP decreases in the presence of glucose, this does not limit mitochondrial ADP uptake, or decrease respiration, due to the very high affinity of the mitochondrial ADP transporter (Diaz-Ruiz et al., 2011; Rodríguez-Enríquez et al., 2001). These collectively reiterate the importance of Pi access and transport to mitochondria in constraining mitochondrial respiration. Indeed, this interpretation can also contextually explain observations from other model systems where mitochondrial Pi transport seems to regulate respiration (Scheibye-Knudsen and Quistorff, 2009; Seifert et al., 2015).

We parsimoniously suggest that Pi access to the mitochondria as a key constraint for mitochondrial repression under high glucose. In a hypothetical scenario, a single-step event in evolution, reducing mitochondrial Pi transporter amounts, and/or increasing glycolytic flux (to deplete cytosolic Pi), will result in whole-scale metabolic rewiring to repress mitochondria. More elaborate regulatory events can easily be imagined as subsequent adaptations to enforce mitochondrial repression. Given the central role played by Pi, something as fundamental as access to Pi will constrain mitochondrial repression. Concurrently, the rapid incorporation of Pi into faster glycolysis can give cells a competitive advantage, while also sequestering Pi in the form of usable ATP. Over the course of evolution, this could conceivably drive other regulatory mechanisms to enforce mitochondrial repression, leading to the currently observed complex regulatory networks and signaling programs observed in the Crabtree effect, and other examples of glucose-dependent mitochondrial repression.

Materials and methods

Statistics and graphing

Unless otherwise indicated, statistical significance for all indicated experiments were calculated using unpaired Student’s t-tests (GraphPad prism 9.0.1). Graphs were plotted using GraphPad prism 9.0.1

Yeast strains, media, and growth conditions

Earlier version: A prototrophic CEN.PK strain of Saccharomyces cerevisiae (WT) (72) was used unless mentioned otherwise. Strains are listed in appendix table S1. Gene deletions, chromosomal C-terminal tagged strains were generated by PCR mediated gene deletion/tagging (73). Mitochondria targeted Mneon strain (Mito-mNeon green) is described in (31). The cox2-62 strain is described in (32). Media compositions, growth conditions and CRISPR-Cas9 based mutagenesis are described in extended methods (SI appendix).

Revised version: A prototrophic CEN.PK strain of Saccharomyces cerevisiae (WT) (J.P. van Dijken et al., 2000) was used unless mentioned otherwise. Strains are listed in appendix table S1. Gene deletions, chromosomal C-terminal tagged strains were generated by PCR mediated gene deletion/tagging (Longtine et al., 1998). Mitochondria targeted Mneon strain (Mito-mNeon green) is described in (Dua et al., 2022). The cox2-62 strain is described in (Bonnefoy et al., 2001). Media compositions, growth conditions and CRISPR-Cas9 based mutagenesis are described in extended methods (SI appendix).

Mitotracker fluorescence

Mitotracker fluorescence was measured using Thermo VarioscanTM LUX multimode plate reader (579/599 excitation/emission). Detailed protocol is described in extended methods. Mitotracker fluorescence were normalized using OD600 of each sample and relative fluorescence intensity calculated.

Protein extraction and Western blotting

Earlier version: Total protein was precipitated, extracted using trichloroacetic acid (TCA) as described earlier (74). Blots were quantified using ImageJ software. Detailed protocol is described in extended methods.

Revised version: Total protein was precipitated, extracted using trichloroacetic acid (TCA) as described earlier (Vengayil et al., 2019). Blots were quantified using ImageJ software. Detailed protocol is described in extended methods.

Basal OCR measurement

The basal OCR was measured using Agilent Seahorse XFe24 analyzer. Basal OCR readings were normalized for cell number (using OD600 of samples) in each well. The detailed methods are described in extended methods.

Mitochondrial volume estimation

Earlier version: High-resolution 3D fluorescence experiments were performed on an inverted confocal laser scanning microscope (Carl Zeiss LSM 780 or Olympus FV3000). For each imaging field of view, sequential z-stacks were acquired for each excitation channel. 488 nm laser excitation for mNeonGreen and 561 nm laser excitation for Mitotracker CMXros dye were used respectively. Images taken were deconvolved and analyzed further in ImageJ software with custom-written routines. Mitochondria segmentation and quantification was done using the Mitochondria Analyzer plugin (75) in Imagej. For visualization, maximum intensity projection of 3D images was used.

Revised version: High-resolution 3D fluorescence experiments were performed on an inverted confocal laser scanning microscope (Carl Zeiss LSM 780 or Olympus FV3000). For each imaging field of view, sequential z-stacks were acquired for each excitation channel. 488 nm laser excitation for mNeonGreen and 561 nm laser excitation for Mitotracker CMXros dye were used respectively. Images taken were deconvolved and analyzed further in ImageJ software with custom-written routines. Mitochondria segmentation and quantification was done using the Mitochondria Analyzer plugin (Chaudhry et al., 2019) in Imagej. For visualization, maximum intensity projection of 3D images was used.

RNA extraction and RT-qPCR

Earlier version: The RNA extraction was done using the hot phenol extraction method as described in (74). The isolated RNA was DNase treated, and used for cDNA synthesis. Superscript III reverse transcriptase enzyme (Invitrogen) was used for cDNA synthesis and RT-qPCR was performed using KAPA SYBR FAST qRT PCR kit (KK4602, KAPA Biosystems). Taf10 was used as a control for normalisation and the fold change in mRNA levels were calculated by 2^-ddct method.

Revised version: The RNA extraction was done using the hot phenol extraction method as described in (Vengayil et al., 2019). The isolated RNA was DNase treated, and used for cDNA synthesis. Superscript III reverse transcriptase enzyme (Invitrogen) was used for cDNA synthesis and RT-qPCR was performed using KAPA SYBR FAST qRT PCR kit (KK4602, KAPA Biosystems). Taf10 was used as a control for normalisation and the fold change in mRNA levels were calculated by 2^-ddct method.

ATP, ethanol and Pi measurements

Earlier version: ATP levels were measured by ATP estimation kit (Thermo Fisher A22066). Ethanol concentration in the medium was estimated using potassium dichromate based assay described in (76) with modifications. Pi was estimated using a malachite green phosphate assay kit (Cayman chemicals, 10009325). Detailed sample collection and assay protocols are described in extended methods.

Revised version: ATP levels were measured by ATP estimation kit (Thermo Fisher A22066). Ethanol concentration in the medium was estimated using potassium dichromate based assay described in (Sriariyanun et al., 2019) with modifications. Pi was estimated using a malachite green phosphate assay kit (Cayman chemicals, 10009325). Detailed sample collection and assay protocols are described in extended methods.

Metabolite extraction and analysis by LC-MS/MS

Earlier version: The steady state levels and relative 13C label incorporation into metabolites were estimated by quantitative LC-MS/MS methods as described in (77). Detailed methodology is extensively described in extended methods. Peak area measurements are listed in supplement file 1.

Revised version: The steady state levels and relative 13C label incorporation into metabolites were estimated by quantitative LC-MS/MS methods as described in (Walvekar et al., 2018). Detailed methodology is extensively described in extended methods. Peak area measurements are listed in supplement file 1.

Mitochondrial isolation

Earlier version: Mitochondria was isolated by immunoprecipitation as described in (78, 79) with modifications. The detailed protocol is described in extended methods (SI appendix). The eluted mitochondria were used in malachite green assay for Pi estimation, boiled with SDS-glycerol buffer for western blots or incubated with mitotracker CMXROS with mitochondria-activation buffer for mitotracker assays.

Revised version: Mitochondria was isolated by immunoprecipitation as described in (Chen et al., 2017; Liao et al., 2018) with modifications. The detailed protocol is described in extended methods (SI appendix). The eluted mitochondria were used in malachite green assay for Pi estimation, boiled with SDS-glycerol buffer for western blots or incubated with mitotracker CMXROS with mitochondria-activation buffer for mitotracker assays.

Added in the revised version: Cytosolic fraction isolation

Added in the revised version: Cytosolic fraction was isolated from spheroplasts by centrifugation as described in detail in S1 appendix. The total protein amounts in the cytosolic fraction was estimated by BCA protein estimation assay and the Pi levels were estimated by malachite green assay.

Supporting information